The TOR signaling pathway regulates vegetative development, aflatoxin biosynthesis, and pathogenicity in Aspergillus flavus

Author Response

The following is the authors’ response to the original reviews.

Response to Reviewer 1 Comments (Public Review):

Point 1: While the authors provided a large amount of data regarding the genes involved in the TOR pathway, it is highly descriptive and mostly confirmative data, as numerous papers have already shown that the TOR pathway plays essential roles in a myriad of biological processes in multiple fungi.

Response 1: Thank you for your comment. The target of rapamycin (TOR) signal pathway plays critical roles in various eukaryotic organisms. However, its specific role in controlling the development and virulence of opportunistic pathogenic fungi like A. flavus has remained unclear. Additionally, the underlying mechanism of the TOR pathway remains elusive in the A. flavus. As such, our study provides a useful contribution, as it is the first to comprehensively investigate the majority of genes in the conserved TOR signaling pathway in A. flavus.

Point 2: The authors seemed to perform a series of parallel studies in several genes involved in the TOR pathway in other fungi. However, their data are not properly interconnected to understand the TOR signaling pathway in this fungal pathogen. The authors frequently drew premature conclusions from basic phenotypic observations. For instance, based on their finding that sch9 mutant showed high calcium stress sensitivity, they concluded that Sch9 is the element of the calcineurin-CrzA pathway. Furthermore, based on their finding that the sch9 mutant show weak rapamycin sensitivity and increased Hog1 phosphorylation, they concluded that Sch9 is involved in TOR and HOG pathways. To make such conclusions, the authors should provide more detailed mechanistic data.

Response 2: Yes, we agree with the reviewer's comment. We have carefully reviewed the manuscript and made necessary revisions to eliminate arbitrary conclusions. For example, we have removed the statement that "Sch9 is the element of the calcineurin-CrzA pathway". Furthermore, we have rephrased our conclusions to better reflect our findings. "these results reflected that Sch9 regulates osmotic stress response via the HOG pathway in A. flavus"(Lines 279-280, page 13). We appreciate the reviewer's input, which has contributed to the clarity and accuracy of our work.

Point 3: In the section "Tor kinase plays important roles in A. flavus", some parts of their data are confusing. The authors said they identified a single Tor kinase ortholog, which is orthologous to S. cerevisiae Tor2. And then, they said failed to obtain a null mutant, but constructed a single copy deletion strain delta Tor1+/Tor2-. What does this mean? Does this mean A. flavus diploid strain? So is this heterozygous TOR/tor mutant? Otherwise, does the haploid A. flavus strain they used contain multiple copies of the TOR gene within its genome? What is the real name of A. flavus Tor kinase (Tor1 or Tor2?). "tor1+/tor2-" is the wrong genetic nomenclature. What is the identity of detalTor1+/Tor2-? Please provide detailed information on how all these mutants were generated. A similar issue was found in the analysis of TapA, which is speculated to be essential (what is the deltaTapA1+/TapA2-?). I couldn't find any detailed information even in Materials and Methods. The authors should provide southern blot data to validate all their mutants.

Response 3: Thank you for your comments. We acknowledge the confusion in our presentation and will ensure that accurate genetic nomenclature is used consistently throughout the paper.

In response to your queries, we have included a section in the Materials and Methods, titled "Detection of tor and tapA genes copy number in strains" (Lines 615-621, page 29), to provide details on how we determined the copy numbers of the tor and tapA genes in the strains. Our findings revealed that both the tor and tapA genes are present in double copies in our strains, which guided our decision to construct single-copy deletion strains using homologous recombination. We have verified these copy numbers using absolute quantification PCR (Table S1).

The use of the abbreviation '+/-' for the single copy knockout strains, such as tor+/- and tapA+/-, is consistent with common fungal literature practice. We apologize for any confusion caused by this nomenclature.

Although we did not employ southern blot data for validation, we conducted PCR and gene sequencing to confirm the mutants. We appreciate your comments to improve the clarity and accuracy of our manuscript.

Point 4: How were the FRB domain deletion mutants constructed? If the FKBP12-rapamycin binding (FRB) domain is specifically deleted in the Tor kinase allele, should it be insensitive and resistant to rapamycin? However, the authors showed that the FRB domain deleted TOR allele was indeed non-functional.

Response 4: We appreciate the reviewer's attention to the construction of the Fkbp12-rapamycin binding (FRB) domain deletion mutants and the discrepancy between the expected and observed results.

For the knockout of the FRB domain, we used the homologous recombination method, but because tor genes are double-copy genes, there are also double copies in the FRB domain. Despite our efforts, we encountered challenges in precisely determining the location of the other copy of the tor gene.

We speculate the common expectation that the deletion of the FRB domain should result in insensitivity and resistance to rapamycin, as it disrupts the binding site for Fkbp-rapamycin. However, we observed that the FRB domain-deleted mutant was more sensitive to rapamycin. This intriguing result suggests that there are additional factors or complexities involved in TOR signaling pathway regulation in A. flavus. We hypothesize that this result is related to the double copy of the tor gene. The reviewer's keen observation and comment have contributed to our efforts to better understand and explain this intriguing result.

Point 5: In Figure 4C, the authors should monitor Hog1 phosphorylation patterns under stressed conditions, such as NaCl treatment, and provide quantitative measurements. Similar issues were found in the western blot analysis of Slt2 (Fig. 8D).

Response 5: We agree with the reviewer that we should monitor Hog1 phosphorylation patterns under stressed conditions. In response to this valuable suggestion, we conducted additional experiments to examine Hog1 phosphorylation patterns under NaCl treatment for 30 minutes. The quantitative measurements of Hog1 phosphorylation levels under stress have been added to Figure 4E in the revised manuscript. Similarly, we have addressed the issue raised regarding Slt2 in Figure 8D.

Point 6: For all the deletion mutants generated in this study, the authors should generate complemented strains to validate their data.

Response 6: We appreciate the reviewer's suggestion to generate complemented strains for all the deletion mutants in our study to validate our data. However, due to the extensive number of genes involved in this research, it is hard to create complemented strains for each individual deletion mutant. As suggested by the reviewer, we have constructed complemented strains for several key deletion mutants, such as ΔsitA-C and Δppg1-C.

Response to Reviewer 1 Comments (Recommendations For The Authors):

Point 1: Overall, this manuscript was very poorly organized and not presented logically. It requires extensive English language editing.

Response 1: We appreciate the reviewer's feedback regarding the organization and language quality of our manuscript. To address these concerns, we have restructured the manuscript to improve its logical flow and coherence. We thank the reviewer for their constructive criticism, which has been instrumental in the manuscript's refinement.

Point 2: The authors did not present their figures in the order of description. For example, the authors suddenly described Figure 9A data in lines 128-130 in the middle of describing Figure 1. Furthermore, Figures 1D and 1F were described earlier than Figures 1B and 1C. In addition, Figure S2 was shown earlier than Figure S1. Please check this throughout the manuscript.

Response 2: We thank the reviewer for their insightful observation. We acknowledge the importance of a logical and coherent figure sequence for reader comprehension. After careful review, we have rearranged the text and images throughout the entire document to enhance the reading experience. The revised manuscript now presents figures in a consistent and logical order, following the sequence of descriptions. We believe this improvement will enhance the overall readability and comprehension of our research.

Point 3: The authors should follow the standard genetic nomenclature rules.

Response 3: Thank you for your suggestion. We have revised our manuscript to ensure that we are following the standard genetic nomenclature rules throughout. This includes the correct naming of genes, proteins, and mutations, as well as the use of appropriate italicization and formatting. We follow the rules: gene symbols are typically composed of three lowercase italicized letters, while protein symbols are not italicized, with an initial capital letter followed by lowercase letters.

Point 4: These are just a few examples. Besides the ones that I mentioned, I found numerous grammatically wrong or awkward sentences throughout the manuscript. So this manuscript requires extensive English proofreading.

Response 4: We apologize for the problem of our manuscript. We have asked an English native speaker to enhance the overall language quality and readability of the text. We believe that these improvements will significantly enhance the manuscript's overall quality and make it more accessible to a broader audience.

Response to Reviewer 2 Comments (Public Review):

Point 1: However, findings have not been deeply explored and conclusions mostly are based on parallel phenotypic observations. In addition, there are some concerns that exist surrounding the conclusions.

Response 1: We are grateful for the suggestion. We conduct additional experiments and analyses to delve more deeply into our findings and ensure a more robust basis for our conclusions.

Response to Reviewer 2 Comments (Recommendations For The Authors):

Point 1: Verification for mutants: a single copy deletion strain ΔTor1+/Tor2(containing one copy of the Tor gene), however, in the table of strain list, it seems like null mutants. There are no further verifications for relative genes' expression and no complementary strains.

A. Flavus ΔTor: Δku70; ΔniaD; ΔTor::pyrG

A. Flavus ΔTapA Δku70; ΔniaD; ΔTapA::pyrG

As described in pp208, "While we failed to obtained a null mutant, we constructed a single copy deletion strain ΔTor1+/Tor2- (containing one copy of the Tor gene) constructed by homologous recombination)"? But the authors think there was only one Tor kinase ortholog (AFLA_044350). It is hard to understand for this mutant What is the evidence to verify phenotypes of the ΔTor1+/Tor2- strain resulted from deletion of Tor2, no detail for how to make ΔTor1+/Tor2- strain.

Response 1: Thank you for your important comments and suggestion. We apologize for the confusion caused by genetic nomenclature. We make the necessary corrections in the table of strain lists to accurately reflect the genotypes of the strains (Table S3).

Multicopy variation of genes has not been explored in detail in fungi, especially in A. flavus, but is a commonly known phenomenon in mammalian genomes[1-2]. In yeast, the presence of two tor genes, tor1 and tor2, whereas in higher eukaryotes such as plants, animals, and filamentous fungi, there is only one tor gene[3-4]. The homology comparison results show that the genome of A. flavus contains only one tor gene. However, the tor gene in A. flavus exhibited varying copy numbers, as was confirmed by absolute quantification PCR at the genome level (Table S1).

In this study, we constructed a single copy deletion strain, tor+/-, through homologous recombination. This strain contains one copy of the tor gene. We provide a more detailed and explicit description of the methods used to detect of the genes copy number in strains (Lines 615-621, page 29). We thank the reviewer for pointing out these important issues.

Point 2: For a point mutant strain TORS1904L, they found that the sensitivity to rapamycin is consistent with the WT strain, it could not tell anything. It should be moved to Suppl.

Response 2: Thanks for your important comments. We acknowledge that these results may not provide significant insights. In response to this suggestion, we delete the data related to the TORS1904L point mutant strain and its sensitivity to rapamycin to ensure that the main manuscript focuses on the most pertinent and informative findings. Corresponding modifications have been made in the revised manuscript.

Point 3: For subtitle "Sch9 is correlate with the HOG and TOR pathways "What is the meaning for "correlate" similarly?

Response 3: Thank you for this comment. We apologize for the unclear wording. To enhance clarity, we revise the subtitle to more explicitly convey this conclusion, for example, "The Sch9 kinase is involved in aflatoxin biosynthesis and the HOG pathway". (Lines 242, page 12).

Point 4:for the ΔTapA 1+/TapA 2- strain (containing one copy of the TapA gene). It should have the complementary strain to verify the specific role of TapA. In FigS1B, ΔTOR and ΔTapA it could not tell TOR gene has been edited. Did you test mRNA of TOR gene?

Response 4: Thanks for your important comments. Due to the large number of genes involved, we did not perform a complementation experiment. However, we used PCR and sequencing to verify the editing of our gene. Additionally, we conducted copy number and mRNA analyses to verify its function. The transcriptional level of the tor gene in the tor+/- mutant was downregulated compared to the level in the wild-type strain (Fig. S6).

Response to Reviewer 3 Comments (Public Review):

Point 1: As for many results, I miss the re-complementation of the created mutants throughout the manuscript. This is standard praxis.

Response 1: Thanks for your suggestions. We acknowledge that re-complementation is a standard practice for validating the effects of gene deletions. However, due to the large number of genes involved in our study, we have performed supplementary experiments on a selection of them, such as ΔsitA-C and Δppg1-C. We are grateful to the reviewer for your understanding of this practical consideration.

Point 2: Fig. 1: cultures were grown for 48 h before measuring the transcript level. The authors show that brlA, abaA, and some sexual regulators are less expressed. In my opinion, this does not allow the conclusion that there is a direct control through rapamycin. Since the colonies grow very slowly in the presence of rapamycin, the authors should add rapamycin and follow gene expression after 15, 30, 60, 90 min. The figure legend needs to be more detailed. Which type of cultures were used, liquid, solid medium? Etc.

Response 2: We deeply appreciate the reviewer’s suggestion. Since we found that there were no significant differences in gene expression changes following shorter treatment times, we extended the treatment duration. We conduct additional experiments to examine the gene expression levels at longer time intervals (3, 6, and 9 h) after the addition of rapamycin (Figure 1H-1J). These time points allow us to capture the dynamic changes in gene expression in response to rapamycin more effectively. Additionally, we enhance the figure legend to provide a more comprehensive description that specifies the type of cultures used in the experiments.

Point 3: Why in chapter one Fig. 9 is already cited? Those data should then be included in Fig. 1 for the general phenotype.

Response 3: Thank you for the suggestion. We have reordered the figures in the updated version of the manuscript to ensure that the data for consistent and clarity.

Point 4: The authors wrote that radial growth and conidiation were gradually reduced with increasing rapamycin concentrations. This is not true. There is no gradient! However, it should be tested if there is a gradient if lower concentrations are used. The current data imply that there is a threshold concentration, so either there is 100 % growth or a reduction to 25 %. This looks strange.

Response 4: Thank you for underlining this deficiency. We agree that a threshold concentration versus a gradient is an important distinction that needs to be clarified. Our results show that the addition of excessive quantities of rapamycin does not increase the inhibition of A. flavus growth. As the concentration of the FK506 drug increases, there is a gradual decrease in the growth and cell production of A. flavus. This phenomenon could potentially be attributed to varying mechanisms of action exhibited by the drugs. Therefore, we have revised these confused sentences. ( Lines 120-121, Page 5)

Point 1: There are many wrong spellings:

Fig. 1. Before washed, before washing; RelaTEtive gene expERSion should read relative gene expression. Sclerotial should be sclerotia. See also Fig. 5 F, H, Fig. 6 E. 6D colon diameter should be colony diameter.

Fig. 4E. The expressED level... should read Expression level..... (also without article) Also in A, F, H.

Fig. 6C. TLC detection of WT.... The authors mean AF detection in extracts of WT..... AF was extracted and analyzed by TLC.....

Labelling of axes in one figure should be uniform.

Response 1: Thank you for your reminder. We apologize for the oversights, and we carefully address and correct all the mentioned spelling issues to ensure the accuracy and clarity of the manuscript.

Point 2: If the authors refer to the genes, I think they should be in small letters and italics, if it is the protein, the first letter should be capitalised tap1 (italics) and Tap1.

Response 2: We appreciate this suggestion. We have carefully checked the entire manuscript and revised follow the standard genetic nomenclature rules. We follow the naming conventions for microbial genes and proteins, where gene symbols are typically composed of three lowercase italicized letters, and protein symbols are not italicized, with an initial capital letter followed by lowercase letters.

Point 3: Very often articles are used where I would not use them.

Response 3: Thanks for your careful checks. We are sorry for our carelessness. Based on your comments, we have made the corrections to make the articles harmonized within the whole manuscript. We value the reviewer's feedback, which will contribute to the overall quality of our writing.

References:

[1] Handsaker R, Van Doren, V, Berman, J. et al. Large multiallelic copy number variations in humans. Nat Genet 47, 296–303 (2015).

[2] Wang Y, Wang S, Nie X. et al. Molecular and structural basis of nucleoside diphosphate kinase-mediated regulation of spore and sclerotia development in the fungus Aspergillus flavus. J Biol Chem. 2019 Aug 16;294(33):12415-12431.

[3] Kim DH, Sarbassov DD, Ali SM, et al. mTOR interacts with raptor to form a nutrient-sensitive complex that signals to the cell growth machinery. Cell. 2002; 110(2): 163-75.

[4] Fu L, Liu Y, Qin G, et al. The TOR-EIN2 axis mediates nuclear signalling to modulate plant growth. Nature. 2021; 591(7849): 288-292.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: The target of the rapamycin (TOR) signaling pathway is highly conserved and important in eukaryotes. It is involved in the regulation of various biological processes. However, systematic studies on this pathway in the genus Aspergillus have not been reported. Here, we identified and characterized nine genes encoding components of the TOR pathway in A. flavus, and investigated their biological, genetic and biochemical functions. The FK506-binding protein FKBP3 and its lysine succinylation are important for aflatoxin production and rapamycin resistance. The Tor kinase plays a central role in the global regulation of growth, spore production, aflatoxin biosynthesis and rapamycin stress. As a major downstream effector molecule of Tor kinase, the Sch9 kinase might regulate the calcium and osmotic stress, AFB1 synthesis of A. flavus by its S_TKc, S_TK_X domains and ATP binding site at K340. We also showed that Sch9 kinase might mediate crosstalk between the TOR and the HOG signaling pathways. TapA and TipA, the other downstream components of Tor kinase, play important roles in regulating mycelial growth and sclerotia formation in A. flavus. The member of the TapA-phosphatase complexes Sit4 and Ppg1 are important for hyphal development, sexual reproduction, sclerotia formation, AFB1 biosynthesis, activation of the CWI and TOR signaling pathways in A. flavus. In addition, the another phosphatase complex Nem1/Spo7 play critical role in vegetative growth, conidiation, aflatoxin and LD biogenesis. This study provide new insights into constructing the regulatory network of the TOR signaling pathway and revealing the molecular mechanism of the pathogenicity in A. flavus.

Revised version: The target of rapamycin (TOR) signaling pathway is highly conserved and plays a crucial role in diverse biological processes in eukaryotes. However, the underlying mechanism of the TOR pathway in A. flavus remains elusive. In this study, we identified and characterized nine genes encoding various components of the TOR pathway in A. flavus, and investigated their biological function. The FK506-binding protein Fkbp3 and its lysine succinylation are important for aflatoxin production and rapamycin resistance. The Tor kinase plays a pivotal role in the regulation of growth, spore production, aflatoxin biosynthesis, and osmotic and rapamycin stress. As a significant downstream effector molecule of the Tor kinase, the Sch9 kinase regulates aflatoxin B1 (AFB1) synthesis, osmotic and calcium stress response in A. flavus, and this regulation is mediated through its S_TKc, S_TK_X domains, and the ATP binding site at K340. We also showed that the Sch9 kinase may have a regulatory impact on the high-osmolarity glycerol (HOG) signaling pathway. TapA and TipA, the other downstream components of the Tor kinase, play significant roles in regulating sclerotia formation and cell wall stress response in A. flavus. The members of the TapA-phosphatase complexes, SitA and Ppg1, are crucial for various biological processes in A. flavus, including vegetative growth, sclerotia formation, AFB1 biosynthesis, and pathogenicity. Furthermore, we showed that SitA and Ppg1 are involved in regulating lipid droplets (LDs) biogenesis and cell wall integrity (CWI) signaling pathway. In addition, another phosphatase complex, Nem1/Spo7, plays critical roles in hyphal development, conidiation, aflatoxin, and lipid droplets biogenesis. This study provides an important insight into the regulatory network of the TOR signaling pathway and the molecular mechanism of aflatoxin biosynthesis in A. flavus.

Introduction

Earlier version: A. flavus is one of the most important phytopathogenic fungi species, which frequently contaminates various plants, causing economic losses in agricultural commodities and posing serious health concerns in developing countries [1]. A. flavus is also an opportunistic human pathogenic fungus that can cause aspergillosis in immune-compromised patients [2]. Aflatoxins (AFs), the toxic and carcinogenic secondary metabolites produced by A. flavus, are the most potent carcinogen present in nature, causing liver cancer in humans and animals [3]. Although most of the enzymatic reactions involved in AF biosynthesis have been elucidated and the genes encoding these enzymes have been isolated [4], the signaling pathway network and specific regulatory processes involved in AF biosynthesis are still unclear. Therefore, exploring the regulatory mechanism of the AF biosynthesis signaling pathway would provide new insights for preventing A. flavus and aflatoxins contamination.

Revised version: Aspergillus flavus is one of the most important phytopathogenic fungi that frequently contaminates various plants, causing economic losses in agricultural commodities and posing serious health concerns in developing countries [1]. A. flavus is also an opportunistic human pathogenic fungus that can cause aspergillosis in immune-compromised patients [2]. Aflatoxins (AFs), the toxic and carcinogenic secondary metabolites produced by A. flavus, are considered the most potent carcinogen present in nature, causing liver cancer in both humans and animals [3]. The aflatoxin biosynthesis is a complex process involving a series of reactions. Although the majority of the enzyme reactions and genes involved in aflatoxin biosynthesis have been elucidated [4], the related signaling pathway network and specific regulatory mechanisms governing aflatoxin biosynthesis remain elusive. Therefore, exploring the regulatory mechanism of the aflatoxin biosynthesis signaling pathway would provide new insights for preventing A. flavus and aflatoxins contamination.

Earlier version: In eukaryotes, the target of the rapamycin (TOR) signaling pathway is highly conserved and plays essential roles in several important biological processes, such as ribosome biosynthesis, cell growth and autophagy [5]. Several studies have shown that the TOR signaling pathway interacts with other signaling pathways (MAPK, CWI), and they crosstalk to form a complex metabolic network to regulate cellular processes [6–9]. The serine/threonine-protein kinase Tor acts as a crucial protein element in the TOR signaling pathway, which interacts with other proteins to form multiprotein complexes. There are two TOR complexes named TORC1 and TORC2 in Saccharomyces cerevisiae [10]. TORC1-mediated signaling controls cell growth by regulating many growth-related processes and is sensitive to rapamycin, while TORC2-mediated signaling is mainly responsible for cytoskeletal remodeling and is insensitive to rapamycin [11]. However, most organisms, including plants, animals and filamentous fungi, have a single Tor kinase [12–13].

Revised version: In eukaryotes, the target of rapamycin (TOR) signaling pathway is highly conserved and plays essential roles in various significant biological processes, including ribosome biosynthesis, cell growth, and autophagy [5]. The serine/threonine-protein kinase Tor serves as a crucial protein component in the TOR signaling pathway, which interacts with other proteins to form multi-protein complexes [6]. In Saccharomyces cerevisiae, TOR exists in two distinct multiprotein functional complexes, namely TORC1 and TORC2 [7]. TORC1-mediated signaling governs cellular growth by modulating various growth-related mechanisms and exhibits sensitivity to rapamycin. Conversely, TORC2-mediated signaling primarily regulates cytoskeletal remodeling but remains unaffected by rapamycin [7]. In yeast, there are two tor genes, tor1 and tor2, whereas in higher eukaryotes such as plants, animals, and filamentous fungi, there is only one tor gene [8–10].

Earlier version: In the fungal kingdom, most studies on the TOR signaling pathway have been conducted in budding yeast S. cerevisiae [14–15], and fission yeast Schizosaccharomyces pombe [6.16-17]. In S. cerevisiae, rapamycin forms a complex with the peptidylprolyl cis/trans isomerase FKBP12, which binds and inhibits Tor kinase. The Tap42 phosphatase complex is the main target of Tor kinase, and the Tap42 protein interacts with the Tip41 protein to jointly regulate phosphatases such as Pp2A and Sit4. Phosphatase Sit4 can dephosphorylate transcription factor GLN3, regulating the nitrogen source metabolism pathway [18–19]. In addition, the other immediate effector of TORC1 is AGC kinase Sch9, which participates in other cellular functions including ribosome biosynthesis [20]. In filamentous fungi, the TOR signaling pathway has been reported in Fusarium graminearum [21–22], Magnaporthe oryzae [23], Phanerochaete chrysosporium [24], Podospora anserine [25], A. nidulans [26], A. fumigatus [51–52], F. fujikuroi [27], Botrytis cinerea [28], F. oxysporum [29]and Mucor circinelloides [30]. In a word, filamentous fungi TOR signaling branch mediates spatial and temporal control of cell growth by activating anabolic processes such as ribosome biosynthesis, lipid droplet biogenesis, carbon and nitrogen metabolism, protein, lipid, and nucleotide synthesis, and repressing catabolic processes such as autophagy and apoptosis.

Revised version: In the fungal kingdom, the TOR signaling pathway has been extensively investigated in three prominent species: the budding yeast S. cerevisiae [11–12], the fission yeast Schizosaccharomyces pombe [13–15], and the human pathogen Candida albicans [16–17]. In S. cerevisiae, macrolide antibiotic rapamycin interacts with the cytosolic protein Fkbp12 to form a complex, and the Fkbp12-rapamycin complex effectively inhibits the activity of the Tor kinase function by binding to the FRB domain [18]. The Tap42 phosphatase complexs serves as the primary target of the Tor kinase. Additionally, the Tap42 protein interacts with the Tip41 protein to collaboratively regulate phosphatases, including PP2A and Sit4 [19–20]. Phosphatase Sit4 can dephosphorylate the transcription factor Gln3, regulating the nitrogen source metabolism pathway [21]. In addition, another immediate effector of TORC1 is the AGC kinase Sch9, which participates in various cellular functions processes, such as ribosome biosynthesis, translation initiation, and entry into G0 phase [22]. The TOR signaling pathway has been documented in various filamentous fungi, including Fusarium graminearum [23–24], Magnaporthe oryzae [25], Phanerochaete chrysosporium [26], Podospora anserine [27], A. nidulans [28], A. fumigatus [29–30], F. fujikuroi [31], Botrytis cinerea [32], F. oxysporum [33], and Mucor circinelloides [34]. However, the TOR signaling pathway has not yet been reported in A. flavus.

Added in the revised version: In eukaryotic cells, the TOR signaling pathway and the Mitogen Activated Protein Kinase (MAPK) signaling pathways play an important role in regulating adaptive responses to extra- and intracellular conditions [35]. Several studies have demonstrated that the TOR signaling pathway interacts with various other signaling pathways, such as MAPK and CWI. These pathways engage in crosstalk, forming a complex metabolic network that regulates cellular processes [36–40]. However, the underlying mechanisms of multiple crosstalks between the TOR and other signaling pathways in A. flavus remain unclear. Therefore, we intended to identify the genes associated with the TOR signaling pathway and elucidate their roles and contributions in regulating vegetative development and aflatoxin biosynthesis in A. flavus. Our results demonstrated that the TOR signaling pathway is involved in multiple cellular processes in A. flavus. Our findings revealed a complex interplay between the TOR pathway and other signaling pathways (such as HOG and CWI) in A. flavus, which play a crucial role in regulating cellular growth and survival under environmental stress conditions.

Removed in the revised version: In yeast and filamentous fungi, the TOR signaling pathway modulates cell growth and vegetative differentiation by sensing changes in external nutrients, energy and stress [31]. While the underlying mechanisms of multiple crosstalks between TOR and the other signaling pathway in A. flavus pathogens are unclear. Therefore, we intented to identify the gene of TOR signaling pathway and demonstrate their roles and contributions in the regulation of vegetative development and aflatoxin biosynthesis in A. flavus. In this study, the results indicated that the TOR signaling pathway plays crucial roles in various cellular processes including vegetative development, aflatoxin biosynthesis and pathogenicity in A. flavus. Our results also indicated a highly complicated cross-regulatory relationship between TOR and other signaling pathways (MAPK, CWI) that maintain cellular growth and survival in response to environmental stress in A. flavus.

Results

Earlier version: Rapamycin inhibits A. flavus growth, sporulation, sclerotia and aflatoxin production

Revised version: Rapamycin exhibits inhibitory effects on the growth, sporulation, sclerotia formation, and aflatoxin production in A. flavus

Added in the revised version: Rapamycin, a secondary metabolite synthesized by Streptomyces hygroscopicus, exhibits efficacy in suppressing specific filamentous fungi. In F. graminearum, rapamycin has a significant inhibitory impact on growth and asexual reproduction [23]. To investigate the effects of rapamycin on vegetative development, sclerotia formation, and secondary metabolism in A. flavus, we conducted an assay to determine the sensitivity of the wild-type strain to rapamycin. We observed that the wild-type strain exhibited a high sensitivity to rapamycin. We found that 100 ng/mL of rapamycin significantly inhibited mycelial growth and conidia formation (Fig.1A, 1D, 1E). More importantly, the wild-type strain exhibited a significant reduction in sclerotia formation and aflatoxin synthesis (Fig.1B, 1C, 1F, 1G). Additionally, rapamycin resulted in a significant decrease in spore germination rate and sparser conidiophores (Fig. S1A, S1B). These results illustrated that the TOR signaling pathway may be involved in regulating the growth, development, and secondary metabolism in A. flavus. Calcofluor white is a chemifluorescent blue dye commonly employed for its nonspecific binding affinity to fungal chitin. Fluorescence microscopy analysis revealed a notable reduction in the blue fluorescence on the cell wall following treatment with rapamycin (Fig.S1C). The intracellular lipid droplets within the hyphae were labeled with BODIPY. We observed that the hyphae treated with rapamycin contained more lipid droplets compared to the control group (Fig. S1D). These findings suggest that the TOR pathway may play important roles in maintaining cell wall integrity and facilitating the biogenesis of lipid droplets.

Removed in the revised version: Rapamycin, a secondary metabolite produced by Streptomyces hygroscopicus, is useful in the inhibition of certain filamentous fungi. In F. graminearum, rapamycin exhibited a strong inhibitory effect on growth and asexual reproduction [21], which also induces lipid droplet (LD) accumulation in F. graminearum and other filamentous fungi [22]. To investigate the effect of rapamycin on vegetative development, secondary metabolism and LD accumulation of A. flavus, we assayed the sensitivity of the WT strain to rapamycin. We found that the A. flavus WT strain was very sensitive to rapamycin, and the mycelial growth, the germination of dormant spores and conidia formation of the WT strain were significantly inhibited on PDA solid plates supplemented with 100 ng/mL rapamycin (Fig. 1A). The radial growth and the spore production of A. flavus WT strain were gradually decreased with the increasing concentration of rapamycin(Fig. 1D,1F). In addition, the rapamycin also showed intense inhibition on sclerotia and AFB1 biosynthesis. As shown in Fig. 1B and 1C, compared to untreated strains, A. flavus was unable to synthesize aflatoxin and form sclerotia in PDA added with 100 ng/mL rapamycin. Furthermore, we observed that A. flavus displayed a dramatic reduction in the rate of spore germination and sparser conidiophores with rapamycin treatment (Fig. S2A,S2B). Calcofluor white is a chemifluorescent blue dye that is nonspecifically used to bind to the chitin of fungi. Fluorescence microscopy showed that the blue fluorescence on the cell wall is significantly diminished when treated with rapamycin (Fig. S2C). The intracellular lipid droplets were stained by BODIPY in the hyphae, and we found that the hyphae treated with rapamycin contained many more LDs (Fig. 9A). These findings suggested that rapamycin-treated hyphae led to the accumulation of LD biogenesis and reduction of chitin content. To explore the potential regulatory function of the TOR pathway to affect the growth, sporulation, sclerotia and aflatoxin biosynthesis of A. flavus, qRT-PCR was used to determine the expression levels of genes related to conidia, sclerotia and aflatoxin biosynthesis. The transcription levels of sporulation and sclerotia genes were both significantly downregulated in rapamycin treatment strains, while the genes related to aflatoxin synthesis were almost not expressed(Fig. 1G-1I), which indicated that the TOR pathway is involved in growth, sporulation, sclerotia and aflatoxin biosynthesis in A. flavus by regulating these genes. All the above results showed that rapamycin affects the growth, sporulation, sclerotia and aflatoxin synthesis of A. flavus, indicating that the TOR pathway plays important roles in vegetative growth and secondary metabolism in A. flavus.

Earlier version: Impacts of rapamycin on growth, sporulation, sclerotia and aflatoxin production of Aspergillus flavus.

Revised version: Impacts of rapamycin on the growth, sporulation, sclerotia formation, and aflatoxin production of Aspergillus flavus

Earlier version: (A) The A. flavus wild-type (WT) strain was incubated on PDA amended with different concentrations of rapamycin at 37 °C for 5 days.

Revised version: (A) Colony morphology of the WT strain cultured on PDA medium amended with different concentrations of rapamycin at 37°C for 5 days.

Earlier version: (B) Colony phenotype of WT strain grown on CM medium amended with different concentrations of rapamycin at 37 °C for 7 days.

Revised version: (B) Colony phenotype of the WT strain grown on CM medium amended with different concentrations of rapamycin at 37°C for 7 days.

Earlier version: (C) TLC assay of AFB1 production by WT strain cultured in YES liquid media at 29°C for 6 days.

Revised version: (C) TLC assay of AFB1 production by the WT strain cultured in YES liquid media amended with different concentrations of rapamycin at 29°C for 6 days.

Earlier version: (D) Statistical analysis of the diameter by the WT strain treated with different concentrations of rapamycin described in (A).

Revised version: (D) Statistical analysis of the colony diameter by the WT strain treated with different concentrations of rapamycin described in (A).

Earlier version: (E) Conidial quantification of the WT strain treated with different concentrations of rapamycinas as mentioned in (A).

Revised version: (E) Conidial quantification of the WT strain treated with different concentrations of rapamycin as mentioned in (A).

Earlier version: (F) Quantitative analysis of sclerotium formation by the WT strain treated with different concentrations of rapamycin described in (C).

Revised version: (F) Quantitative analysis of sclerotium formation by the WT strain treated with different concentrations of rapamycin described in (B).

Added in the revised version: (G) Relative quantification of AFB1 production by the WT strain treated with different concentrations of rapamycin as mentioned in (C)

Earlier version: (G) Relative transcript levels of conidia-related genes in WT strain treated with 100 ng/mL rapamycin cultured for 48 h.

Revised version: (H) Relative expression levels of conidia-related genes in WT strain treated with 100 ng/mL rapamycin after 3, 6, 9 h.

Earlier version: (H) Relative transcript levels of sclerotia-related genes in WT strain treated with 100 ng/mL rapamycin cultured for 48 h.

Revised version: (I) Relative expression levels of sclerotia-related genes in WT strain treated with 100 ng/mL rapamycin after 3, 6, 9 h.

Earlier version: (I) Rative transcript levels of aflatoxin biosynthesis regulatory and structural genes in WT strain treated with 100 ng/mL rapamycin cultured for 48 h.

Revised version: (J) Relative expression levels of aflatoxin biosynthesis regulatory and structural genes in WT strain treated with 100 ng/mL rapamycin after 3, 6, 9 h.

Added in the revised version: To investigate the potential regulatory role of the TOR pathway in modulating the sporulation, sclerotia formation, and aflatoxin biosynthesis in A. flavus, we examined the expression levels of conidia-related, sclerotia formation-related, and aflatoxin-related genes after 3, 6, and 9 hours of treatment with rapamycin. To determine the effects of rapamycin on conidiation, qRT-PCR was conducted to assess the transcript levels of two conidia-related genes, abaA and brlA. We also conducted qRT-PCR to evaluate the expression levels of three regulators (nsdC, nsdD, and sclR) involved in sclerotia development, and genes involved in AFB1 biosynthesis, such as the regulatory gene aflS, as well as several structural genes (aflC and aflQ). The strains treated with rapamycin exhibited a significant reduction in the expression levels of genes related to sporulation, sclerotia formation, and aflatoxin synthesis (Fig. 1H, 1I, 1J). Overall, these results indicated that the TOR pathway plays a crucial role in regulating the growth, sporulation, sclerotia formation, and aflatoxin biosynthesis in A. flavus by modulating the expression of these genes.

Earlier version: Effect of FKBP3 and its lysine succinylation on AFB1 biosynthesis and rapamycin resistance in A. flavus

Revised version: The impact of Fkbp3 and its lysine succinylation on AFB1 biosynthesis and rapamycin resistance in A. flavus

Earlier version: FK506-binding proteins (FKBPs) are highly conserved immunophilin family and its members involved in various cellular functions including protein folding, immunosuppression, signaling transduction and transcription[33]. FKBPs act as molecular switches, which elicit function through binding and altering the conformation of target proteins. In S. cerevisiae, FKBP12 and rapamycin form a complex that negatively regulates the Tor kinase activity[33]. According to the protein sequence of peptidylprolyl isomerase fprA from A. nidulans. Four genes encoding putative proteins with FK506 binding domain, named FKBP1(AFLA_126880), FKBP2(AFLA_087480), FKBP3 (AFLA_010910) and FKBP4(AFLA_128200) were identified from the A. flavus genome. These FKBPs contain a single FKBP-type peptidylprolyl cis-trans isomerase conserved domain. Disruption strains for each FKBP gene were generated by homologous recombination using the fusion PCR approach. All FKBP disruption strains were confirmed by PCR and sequencing analysis. Morphological analyses showed that mycelial growth and conidia formation of all the FKBP mutants were similar to that of the WT strain on the PDA medium, suggesting that FKBPs had no significant effects on the vegetative development (Fig. 2A). However, the number of sclerotia was significantly reduced in all FKBP mutants compared to that of the WT strain (Fig. S3), indicating that FKBPs were essential for sclerotia formation. We also assessed AFB1 production via thin layer chromatography (TLC), and the results showed that the deletion of FKBP3 led to an obvious decrease in AFB1 production compared to that of the WT strain (Fig. 2B). Taken together, these data indicated that FKBPs are involved in sclerotia and aflatoxin biosynthesis in A. flavus.

Revised version: FK506-binding proteins (Fkbps) belong to a highly conserved immunophilin family, which are involved in various cellular functions, such as protein folding, immunosuppression, signaling transduction, and transcription [41]. Fkbps function as molecular switches by binding to target proteins and inducing conformational changes [41]. In S. cerevisiae, Fkbp12 and rapamycin interact to form a complex that exerts a negative regulatory effect on the activity of the Tor kinase [42]. The fkbp12 gene homologue fprA in A. nidulans shares high levels of homology with that from S. cerevisiae [28]. According to the protein sequence of FprA from A. nidulans, four genes encoding putative proteins with FK506 binding domain, named fkbp1 (AFLA_126880), fkbp2 (AFLA_087480), fkbp3 (AFLA_010910), and fkbp4 (AFLA_128200), have been identified from the A. flavus genome. We generated four fkbp knockout strains using homologous recombination, and all fkbp disruption strains were confirmed by PCR and sequencing analysis (Fig.S2). Phenotypic analyses showed that the mycelial growth and conidia formation of all the fkbp mutants were comparable to the wild-type strain on the PDA medium (Fig. 2A), suggesting that Fkbps have minimal influence on the process of vegetative development. We evaluated AFB1 synthesis utilizing thin layer chromatography (TLC), and the findings demonstrated a noteworthy reduction in AFB1 synthesis upon deletion of Fkbp3, in comparison to the wild-type strain (Fig.2B, 2D). Additionally, the quantity of sclerotia exhibited a significant decrease in all fkbp mutants compared to the wild-type strain (Fig.S3A, S3B). The aforementioned data indicated that Fkbp3 plays a crucial role in sclerotia formation and aflatoxin biosynthesis in A. flavus.

Disruption of fkbp3 significantly increases the resistance of A. flavus to rapamycin and FK506

Earlier version: (A) Phenotype of A. flavus WT and all mutant strains (ΔFKBP1, ΔFKBP2, ΔFKBP3 and ΔFKBP4) grown on PDA amended with 100 ng/mL rapamycin at 37°C for 5 days.

Revised version: (A) Phenotype of the WT and all mutant strains (Δfkbp1, Δfkbp2, Δfkbp3, and Δfkbp4) grown on PDA amended with rapamycin and FK506 at 37°C for 5 days.

Earlier version: (B) TLC assay of AFB1 production by the WT and all mutant strains (ΔFKBP1, ΔFKBP2, ΔFKBP3 and ΔFKBP4) cultured in YES liquid medium at 29°C for 6 days.

Revised version: (B) TLC analysis of AFB1 production by the WT and all mutant strains (Δfkbp1, Δfkbp2, Δfkbp3, and Δfkbp4) cultured in YES liquid medium at 29°C for 6 days.

Earlier version: (C) The growth inhibition rate of the WT and all mutant strains (ΔFKBP1, ΔFKBP2, ΔFKBP3 and ΔFKBP4) under rapamycin stress.

Revised version: (C) The growth inhibition rate of the WT and all mutant strains (Δfkbp1, Δfkbp2, Δfkbp3, and Δfkbp4) under rapamycin and FK506 stress.

Earlier version: (D) AFB1 quantitative analysis of the WT and all mutant strains (ΔFKBP1, ΔFKBP2, ΔFKBP3 and ΔFKBP4) described in (B).

Revised version: (D) AFB1 quantitative analysis of the WT and all mutant strains (Δfkbp1, Δfkbp2, Δfkbp3, and Δfkbp4) as described in (B).

Earlier version: (E) Phenotype of WT and all mutant strains (ΔFKBP3, K5A, K19A, K40A, K42A, K55A and K65A) grown on PDA amended with 100 ng/mL rapamycin at 37°C for 5 days.

Revised version: (E) Phenotype of the WT and all mutant strains (Δfkbp3, K5A, K19A, K40A, K42A, K55A, and K65A) grown on PDA amended with 100 ng/mL rapamycin at 37°C for 5 days.

Earlier version: (F) The growth inhibition rate of WT and all mutant strains (ΔFKBP3, K5A, K19A, K40A, K42A, K55A and K65A) under rapamycin stress.

Revised version: (F) The growth inhibition rate of the WT and all mutant strains (Δfkbp3, K5A, K19A, K40A, K42A, K55A, and K65A) under rapamycin stress.

(G) TLC assay of AFB1 production by the WT and all mutant strains (Δfkbp3, K5A, K19A, K40A, K42A, K55A, and K65A) cultured in YES liquid medium at 29°C for 6 days.

Earlier version: (H) Relative quantification of AFB1 production in WT and all mutant strains (ΔFKBP3, K5A, K19A, K40A, K42A, K55A and K65A) as mentioned in (G).

Revised version: (H) Relative quantification of AFB1 production in the WT and all mutant strains (Δfkbp3, K5A, K19A, K40A, K42A, K55A, and K65A) as mentioned in (G).

Earlier version: We found that the A. flavus WT strain was very sensitive to rapamycin. To determine which FKBP is involved in rapamycin sensitivity, all mutants were cultured in the PDA solid plates supplemented with 100 ng/mL rapamycin. The FKBP3 disrupted strain showed high resistance to rapamycin, but the FKBP1, FKBP2 and FKBP4 deletion mutants showed no change in resistance to rapamycin (Fig. 2A). These results demonstrated that FKBP3, rather than FKBP1, FKBP2 and FKBP4, is a major target for rapamycin in A. flavus. For FK506-sensitivity analysis, we found that mycelial growth and conidia formation of A. flavus were severely inhibited on PDA added with 10 ng/mL FK506, and the deletion of FKBP3 resulted in increased resistance of A. flavus to FK506 compared to WT strain (Fig. 2A), indicating that the FKBP3 was mainly involved in the regulation of FK506 resistance.

Revised version: To determine the specific Fkbp responsible for rapamycin sensitivity, all mutant strains were cultured on PDA solid plates amended with 100 ng/mL rapamycin. The strain with disrupted fkbp3 showed significant resistance to rapamycin, while the deletion mutants of fkbp1, fkbp2, and fkbp4 did not exhibit any alteration in resistance to rapamycin (Fig.2A, 2C). These results demonstrated that Fkbp3 serves as the principal target for rapamycin in A. flavus. In the FK506-sensitivity analysis, we observed a significant suppression of mycelial growth and conidia formation in A. flavus with the addition of 10 ng/mL FK506 to the PDA medium. Additionally, the deletion of fkbp3 led to the heightened resistance to FK506 compared to the wild-type strain (Fig.2A, 2C). These findings suggested that Fkbp3 is mainly involved in the modulation of FK506 resistance in A. flavus.

Earlier version: In a previous study, we accomplished a global succinylome analysis of A. flavus [60]. Based on our succinylome result, we identified six reliable succinylation sites (K5, K19, K40, K42, K55, K65) on FKBP3. Further evolutionary conservation analysis revealed that lysine at 19 of FKBP3 was highly conserved in the A. flavus orthologs, suggesting that this residue of FKBP3 may be important for an evolutionarily conserved function, which may play a crucial role in regulating aflatoxin biosynthesis in A. flavus. To these succinylation sites, the site-directed mutations were constructed according to the homologous recombination strategy. Compared with the A. flavus WT strain, all point mutants exhibited similar phenotype in vegetative growth and conidiation. Interestingly, we observed that K19A displayed increased resistance to rapamycin, and AFB1 production was significantly reduced when compared to the WT strain (Fig. 3E, 3G), suggesting an essential role of K19 succinylation of FKBP3 in rapamycin resistance and aflatoxin biosynthesis. In addition, other point mutants, including K40A, K42A and K55A, produced less AFB1 compared to the WT strain (Fig. 3G), illustrating that the succinylation of FKBP3 on others site are involved in AF biosynthesis in A. flavus. However, the K19 site likely plays a more prominent role than others sites in FKBP3 succinylation.

Revised version: Previous studies have shown that lysine succinylation plays a significant role in aflatoxin production and pathogenicity in A. flavus [43]. Based on the analysis of our succinylome data, we have identified six reliable succinylation sites (K5, K19, K40, K42, K55, and K65) on the Fkbp3 protein. To identify the function of Fkbp3 succinylation sites, site-directed mutations were constructed according to the homologous recombination strategy. Compared to the wild-type strain, all point mutants exhibited a comparable phenotype in vegetative growth and conidiation. Interestingly, we observed that the K19A mutation resulted in increased resistance to rapamycin (Fig. 2E, 2F). Additionally, the K19A strain exhibited a significant reduction in the production of AFB1 compared to the wild-type strain (Fig. 2G, 2H). These findings suggested that the succinylation of Fkbp3 at K19 plays a crucial role in rapamycin resistance and aflatoxin biosynthesis. Furthermore, we found that various point mutants, including K40A, K42A, and K55A, exhibited reduced levels of AFB1 production in comparison to the wild-type strain (Fig. 2G, 2H), suggesting that the succinylation of Fkbp3 at other sites also contributes to AFB1 biosynthesis in A. flavus. The above results indicated that the K19 site likely plays a more prominent role in Fkbp3 succinylation compared to other sites.

Earlier version: Tor kinase plays important roles in A. flavus.

Revised version: The Tor kinase plays a critical role in A. flavus

Removed in the revised version: (A) The structure diagram of Tor kinase.

Removed in the revised version: (B) Morphology of conidiophores of A. flavus WT, ΔTor1+/Tor2−, TORΔFRB and TORS1904L strains grown on PDA medium at 37°C for 5 days as observed by microscope.

Removed in the revised version: (C) Phenotype of WT, ΔTor1+/Tor2−, TORΔFRB and TORS1904L strains grown on PDA amended with 100 ng/mL rapamycin at 37°C for 5 days.

Removed in the revised version: (D) The growth inhibition rate of WT, ΔTor1+/Tor2−, TORΔFRB and TORS1904L strains under rapamycin stress.

Removed in the revised version: (E) Relative transcript levels of conidia-related genes in WT, ΔTor1+/Tor2−, TORΔFRB strains.

Removed in the revised version: (F) TLC assay of AFB1 production by the WT, ΔTor1+/Tor2−, TORΔFRB and TORS1904L strains cultured in YES liquid medium at 29°C for 6 days.

Removed in the revised version: (G) Quantification analysis of AFB1 as in (F).

Removed in the revised version: (H) Relative transcript levels of aflatoxin biosynthesis regulatory and structural genes in WT and ΔTor1+/Tor2−strains.

Removed in the revised version: Tor kinase plays important roles in A. flavus

Earlier version: In S. cerevisiae, the FKBP12 and rapamycin forms a gain-of-function complex interacts specifically with the TOR proteins. The Tor kinases can bind to various proteins to form complexes such as TORC1 and TORC2. TORC1 can sense signals such as nutrients, growth factors, and stress in the environment, and participate in the regulation of signaling pathways such as gene transcription, protein translation, ribosome synthesis and autophagy[48]. To identify the ortholog of the S. cerevisiae TOR in A. flavus, the S. cerevisiae TOR protein was used as search queries by BLAST (https://blast.ncbi.nlm.nih.gov/Blast.cgi) in the A. flavus genome databasein. The A. flavus genome has a single Tor kinase ortholog (AFLA_044350) encoding a protein with 48.91% similarity to S. cerevisiae TOR2. To determine the function of Tor kinase in A. flavus, we targeted Tor for gene deletion. While we failed to obtained a null mutant, we constructed a single copy deletion strain ΔTor1+/Tor2− (containing one copy of the Tor gene) constructed by homologous recombination. Compared with the WT strain, the ΔTor1+/Tor2− strain grew significantly slower on PDA (Fig. 3C), and produced fewer conidia and sclerotia than the WT strain (Fig. 3B,S4). The results from TLC showed that AFB1 production of the ΔTor1+/Tor2− strain was dramatically decreased compared to that of WT strain (Fig. 3F). Furthermore, The qRT-PCR results showed that the expression levels of the sexual sporulation and aflatoxin synthesis genes were both downregulated in the ΔTor1+/Tor2− mutant compared to WT(Fig. 3E,3H). Interestingly, the ΔTor1+/Tor2− strain cannot survive on PDA amended with 100 ng/mL rapamycin (Fig. 3C). In addition, the ΔTor1+/Tor2− strain was more sensitive to osmotic stress and cell wall inhibitor, including NaCl, KCl, CFW and CR (Fig.S5). These findings suggested that TOR is involved in vegetative growth, asexual development, sclerotia formation, environmental stress and AF biosynthesis in A. flavus.

Revised version: In S. cerevisiae, the Tor kinase can interact with a diverse range of proteins, resulting in the formation of complex structures known as TORC1 and TORC2 [44]. TORC1 can detect various signals, including nutrients, growth factors, and environmental stress. It plays a crucial role in regulating cellular processes such as gene transcription, protein translation, ribosome synthesis, and autophagy [45]. To identify the ortholog of the Tor protein in A. flavus, the S. cerevisiae Tor protein was utilized as search queries in the A. flavus genome database using the Basic Local Alignment Search Tool (https://blast.ncbi.nlm.nih.gov/Blast.cgi). The genome of A. flavus contains a solitary ortholog of the tor gene (AFLA_044350), which encodes a protein exhibiting 48.91% similarity to S. cerevisiae Tor2. We used the absolute quantitative PCR technique to identify the copy number of the tor gene at the genome level in A. flavus (Table S1). To elucidate the function of the Tor kinase in A. flavus, we employed gene deletion techniques to target the tor gene. While our attempt to obtain a null mutant was unsuccessful, we constructed a single copy deletion strain, tor+/−, through homologous recombination, which contains one copy of the tor gene (Table S1). Compared to the wild-type strain, the tor+/− strain exhibited a significantly slower growth rate on PDA (Fig. 3C). Additionally, the tor+/− strain produced fewer conidia and sclerotia compared to the wild-type strain (Fig. 3B, S4). The results from thin-layer chromatography (TLC) revealed a noteworthy decrease in AFB1 synthesis in the tor+/− strain compared to the wild-type strain (Fig. 3D, 3F). Interestingly, the tor+/− mutant exhibited a notable increase in sensitivity to rapamycin and decreased sensitivity to osmotic stress (Fig. 3C, 3E, S5). Furthermore, the qRT-PCR results demonstrated a significant downregulation in the expression levels of the genes associated with sexual sporulation and aflatoxin synthesis in the tor+/− mutant compared to the wild-type strain (Fig. S6). The transcriptional level of the tor gene in the tor+/− mutant was downregulated compared to the level in the wild-type strain (Fig. S6). These findings suggested that the Tor kinase plays a crucial role in various biological processes in A. flavus, including vegetative growth, asexual development, sclerotia formation, response to environmental stress, and biosynthesis of aflatoxins.

Added in the revised version: The Tor kinase plays a vital role in A. flavus

Removed in the revised version: TOR is a conserved serine/threonine kinase, which possesses HEAT repeats at the N-terminal region, and the DUF3385, the FAT, Fkbp12-rapamycin binding (FRB), kinase catalytic and FATC domains at the C-terminal region (Fig. 3A).To assess the role of the FRB domain inside Tor kinase, the FRB domain was deleted with the method of homologous recombination. Compared with the ΔTor1+/Tor2− strain, the TORΔFRB single-copy deletion mutants exhibited similar defects in vegetative growth, conidiation, sclerotia formation, cell wall stress and rapamycin sensitivity (Fig. 3B,3C,S4,S5). In F. graminearum, the S1866L mutation in the FgTor kinase confers rapamycin resistance [21]. To investigate the function of this conserved site, the TORS1904L mutant was constructed by homologous recombination, but site-specific mutation did not impair rapamycin resistance (Fig. 3C). Based on these observations, we speculated that the FRB domain plays critical roles in the regulation of Tor kinase on fungal development and rapamycin resistance in A. flavus.

Removed in the revised version: Sch9 is involved in regulating sclerotia formation and aflatoxin biosynthesis

Removed in the revised version: Sch9, a Ser/Thr kinase of the AGC family, is required for TORC1 to properly regulate ribosome biogenesis, translation initiation, and entry into the G0 phase in S. cerevisiae, and TORC1-dependent phosphorylation is required for Sch9 activity [49]. The A. flavus putative Sch9 (AFLA_044350) gene encode a 706 aminoacid protein with 80.31% identity to the A. nidulans Sch9. To investigate the function of Sch9 in A. flavus, we attempted to generate Sch9 deletion mutant using a homologous recombination strategy. The ΔSch9 strain showed normal vegetative growth and conidiation similar to the WT strain (Fig. 4B). TLC assay and quantitative analysis showed a significantly decreased aflatoxin production in ΔSch9 strain compared to WT strain (Fig. 4C). These results showed that Sch9 regulates aflatoxin biosynthesis in A. flavus.

Removed in the revised version: Sch9 is correlated with the HOG and TOR pathways.

Earlier version: (A) The structure diagram of Sch9 kinase.

Revised version: (A) The structure diagram of the Tor kinase.

Added in the revised version: (B) Morphology of conidiophores in the WT, tor+/-, and torΔFRBstrains observed by microscope.

Earlier version: (B) Phenotype of A. flavus WT, ΔSch9, Sch9ΔC2, Sch9ΔS_TKc, Sch9ΔS_TK_X and Sch9K340A strains grown on PDA amended with 100 ng/mL rapamycin at 37°C for 5 days.

Revised version: (C) Phenotype of the WT, tor+/-, and torΔFRBstrains grown on PDA amended with 100 ng/mL rapamycin at 37°C for 5 days.

Earlier version: (C) TLC assay of AFB1 production by the WT, ΔSch9, Sch9ΔC2, Sch9ΔS_TKc,Sch9ΔS_TK_Xand Sch9K340A strains cultured in YES liquid medium at 29°C for 6 days.

Revised version: (D) TLC assay of AFB1 production by the WT, tor+/−, and torΔFRBstrains cultured in YES liquid medium at 29°C for 6 days.

Removed in the revised version: (D)Colonies of WT, ΔSch9, Sch9ΔC2, Sch9ΔS_TKc, Sch9ΔS_TK_Xand Sch9K340A strains on PDA media amended with 1 M NaCl and 1 M KCl for 3 days.

Removed in the revised version: (E) The expressed levels of phospho-Hog1 were detected in WT and ΔSch9 strains.

Earlier version: (F) The growth inhibition rate of WT, ΔSch9, Sch9ΔC2, Sch9ΔS_TKc, Sch9ΔS_TK_X and Sch9K340A strains under rapamycin stress.

Revised version: (E) The growth inhibition rate of the WT, tor+/-, and torΔFRBstrains under rapamycin stress.

Added in the revised version: (F) AFB1 quantitative analysis of the WT, tor+/-, and torΔFRBstrains described in (D).

Added in the revised version: The Tor protein is characterized by the presence of various conserved domains, including the DUF3385, FAT, FKBP12-rapamycin binding (FRB), kinase catalytic, and FATC domains, which are located at the N-terminus (Fig. 3A). To investigate the functional significance of the FRB domain within the Tor kinase, the FRB domain was eliminated using homologous recombination. Compared to the tor+/− strain, the torΔFRB mutant displayed comparable deficiencies in vegetative growth, conidiation, sclerotia formation, osmotic stress, and rapamycin sensitivity (Fig. 3B, 3C, S4, S5). Based on these observations, we speculated that the FRB domain plays a pivotal role in the modulation of the Tor kinase activity, thereby influencing fungal development and rapamycin resistance in A. flavus.

Added in the revised version: The Sch9 kinase participates in aflatoxin biosynthesis and the HOG pathway

Added in the revised version: (A) The structure diagram of the Sch9 kinase.

Added in the revised version: (B) Phenotype of the WT, Δsch9, sch9ΔC2, sch9ΔS_TKc, sch9ΔS_TK_X, and sch9K340A strains grown on PDA medium amended with 100 ng/mL rapamycin at 37°C for 5 days.

Added in the revised version: (C) TLC assay of AFB1 production from the WT, Δsch9, sch9ΔC2, sch9ΔS_TKc, sch9ΔS_TK_X, and sch9K340A strains cultured in YES liquid medium at 29°C for 6 days.

Added in the revised version: (D) The phenotypes of the WT, Δsch9, sch9ΔC2, sch9ΔS_TKc, sch9ΔS_TK_X, and sch9K340Astrains on YGT media amended with 1 M NaCl and 1 M KCl for 3 days.

Added in the revised version: (E) The phosphorylation levels of Hog1 in the WT and Δsch9 strains were determined with or without osmotic stress.

Added in the revised version: (F) The growth inhibition rate of the WT, Δsch9, sch9ΔC2, sch9ΔS_TKc, sch9ΔS_TK_X, and sch9K340A strains under rapamycin stress.

(G) Quantitative analysis of AFB1 as shown in (C).

Earlier version: (H) The growth promotion rate of WT, ΔSch9, Sch9ΔC2, Sch9ΔS_TKc, Sch9ΔS_TK_X and Sch9K340A strains under osmotic stress.

Revised version: (H) The growth inhibition rate of the WT, Δsch9, sch9ΔC2, sch9ΔS_TKc, sch9ΔS_TK_X, and sch9K340A strains under osmotic stress.

Removed in the revised version: Sch9 is correlate with the HOG and TOR pathways

Removed in the revised version: In yeast, Sch9 is required for gene expression regulated by Sko1, which is directly targeted by the Hog1 MAP kinase, and Sch9 interacts in vitro with both Sko1 and Hog1[53]. To characterize the roles of Sch9 in stress response, we examined the sensitivity of ΔSch9 to various stresses including osmotic stress and calcium stress. Compared with the WT strain, ΔSch9 significantly increased sensitivity to osmotic stress generated by NaCl and KCl (Fig. 4C). In addition, the ΔSch9 strain was more sensitive to the high concentrations of calcium stress(Fig. S5), these results indicated that the Sch9 as a element for the calcineurin-CrzA pathway in response to calcium stress. In A. flavus, the Sch9 kinase contain protein kinase C conserved region 2 (C2), protein kinase (S_TKc), AGC-kinase C-terminal (S_TK_X) domain and a ATP binding site (Fig. 4A). To explore the role of these domains and ATP binding site in Sch9, both domains deletion strains and Sch9K340A mutant were constructed. Compared with the ΔSch9 mutant, the S_TKc deletion strain (Sch9ΔS_TKc), S_TK_X deletion strain (Sch9ΔS_TK_X) and Sch9K340A mutant exhibited similar defects in osmotic stress and AFB1 production(Fig.4B-4D). Therefore, we speculate that protein kinase, kinase C-terminal domain and the ATP binding site at K340 play important roles in the regulation of Sch9 on osmotic and rapamycin resistance in A. flavus.

Removed in the revised version: To determine if Sch9 was involved in the HOG and TOR pathways in A. flavus, we examined the sensitivity of ΔSch9 to rapamycin. Compared with the WT strain, the ΔSch9 strain was slightly more sensitive to rapamycin(Fig. 4B). The MAKP Hog1 is a element of the high osmolarity glycerol (HOG) pathway, so we further examined the phosphorylation level of Hog1. Western blot analysis showed that the phosphorylation level of Hog1 increased significantly in the ΔSch9 strain (Fig. 4E), indicating that Sch9 influences the Hog1 in response to osmotic stress. Taken together, these results reflected that Sch9 serves as a mediator of the TOR and HOG pathways and regulates aflatoxin biosynthesis and multiple stress responses in A. flavus.

Removed in the revised version: TapA plays a major role in the growth, sporulation, sclerotia and aflatoxin production

Removed in the revised version: Tap42, a phosphatase 2A-associating protein, is the direct target of the Tor kinase in yeast. Recent results suggested that Tap42 mediates many of the Tor functions in yeast, especially those involved in transcriptional modulation [55]. The A. flavus genome has a Tap42 homolog named TapA (AFLA_092770), which predicted to encode an 355 amino-acid protein shares 31.9% identity with S. cerevisiae Tap42. To further determine the function of the TOR signaling pathway regulator TapA in A. flavus, we targeted TapA for gene deletion. We failed to obtain a null mutant, so we generated the ΔTapA 1+/TapA 2− strain (containing one copy of the TapA gene). Compared with the WT strain, the ΔTapA 1+/TapA 2− strain led to reduced hyphal growth and conidiation (Fig. 5A.5D), indicating that TapA is important for vegetative growth and conidiation. In addition, the ΔTapA 1+/TapA 2− strain was unable to produce any sclerotia (Fig. 5B), and almost lost its capability in AFB1 production compared to the WT strain, suggesting that TapA is essential for sclerotia formation and biosynthesis of AFB1 (Fig. 5C). The rapamycin sensitivity test showed that the ΔTapA 1+/TapA 2− strain exhibited decreased sensitivity to rapamycin (Fig. 5A). Furthermore, the expression of nsdC, nsdD and sclR were significantly decreased in ΔTapA1+/TapA 2− strain compared to those in the WT(Fig. 5F). These findings revealed that the TapA plays an important positive role in growth, sporulation, sclerotia and AF biosynthesis in A. flavus.

Earlier version: TapA and TipA regulate growth, sclerotia and aflatoxin biosynthesis in A. flavus.

Revised version: TapA and TipA regulate sclerotia formation in A. flavus

Earlier version: (A) Phenotype of the A. flavus WT, ΔTapA1+/TapA2− and ΔTipA strains grown on PDA amended with 100 ng/mL rapamycin at 37°C for 5 days.

Revised version: (A) Phenotype of the WT, tapA +/−, and ΔtipA strains grown on PDA amended with 100 ng/mL rapamycin at 37°C for 5 days.

Earlier version: (B) Colony phenotype of the WT, ΔTapA1+/TapA2−and ΔTipA strains grown on CM medium at 37°C for 7 days.

Revised version: (B) Phenotypic characterization of the WT, tapA +/−, and ΔtipA strains grown on CM medium at 37°C for 7 days.

Added in the revised version: (C) Morphology of conidiophores in the WT, tapA +/−, and ΔtipA strains observed by microscope.

Earlier version: (C) TLC assay of AFB1 production by the WT, ΔTapA1+/TapA2− and ΔTipA strains cultured in YES liquid medium at 29°C for 6 days.

Revised version: (D) TLC assay of AFB1 production by the WT, tapA +/−, and ΔtipA strains cultured in YES liquid medium at 29°C for 6 days.

Added in the revised version: (E) The growth inhibition rate of the WT, tapA +/−, and ΔtipA strains under rapamycin stress.

Removed in the revised version: (D) Fluorescent images of hyphal cells for the WT, ΔTapA1+/TapA2− and ΔTipA strains stained with calcofluor white (CFW).

Earlier version: (E) Amounts of sclerotia produced by the WT, ΔTapA1+/TapA2− and ΔTipA strains.

Revised version: (F) Amount of sclerotia produced by the WT, tapA +/−, and ΔtipA strains.

Added in the revised version: (G) AFB1 quantitative analysis of the WT, tapA +/−, and ΔtipA strains described in (D).

Added in the revised version: The Sch9 kinase is involved in aflatoxin biosynthesis and the HOG pathway

Added in the revised version: Sch9, a Ser/Thr kinase belonging to the AGC family, is directly phosphorylated by TORC1. The phosphorylation of Sch9 is inhibited by rapamycin, as well as by carbon or nitrogen starvation [22]. The putative sch9 gene (AFLA_127440) in A. flavus encodes a protein consisting of 706 amino acids, exhibiting 80.31% identity to the Sch9 protein in A. nidulans. To investigate the role of Sch9 in A. flavus, we attempted to create a sch9 deletion mutant using a homologous recombination strategy. The Δsch9 strain exhibited typical vegetative growth and conidiation similar to the wild-type strain (Fig. 4B). TLC assay and quantitative analysis showed a significantly decreased aflatoxin production in the Δsch9 strain compared to the wild-type strain (Fig. 4C, 4G). These results indicated that Sch9 regulates aflatoxin biosynthesis in A. flavus.

Added in the revised version: In yeast, Sch9 plays an important role in the transcriptional activation of osmostress inducible genes [46]. To characterize the roles of Sch9 in stress response, we examined the sensitivity of the sch9 deletion strain to various stresses including osmotic stress and calcium stress. Compared with the wild-type strain, the Δsch9 mutant significantly increased sensitivity to calcium stress (Fig. S7), and decreased sensitivity to osmotic stress induced by NaCl and KCl (Fig.4D, 4H). In A. flavus, the Sch9 kinase consists of several domains, including the protein kinase C conserved region 2 (C2), protein kinase (S_TKc), AGC-kinase C-terminal (S_TK_X) domain, and an ATP binding site (Fig. 4A). To explore the role of these domains and the ATP binding site in Sch9 kinase, we constructed deletion strains for both domains and a sch9K340A mutant. The results revealed that the sch9ΔS_TKc strain, sch9ΔS_TK_X strain, and sch9K340A mutant exhibited similar phenotypes in osmotic stress and AFB1 production as the Δsch9 mutant (Fig.4B, 4C, 4D). Therefore, we speculate that the protein kinase domain, kinase C-terminal domain, and the ATP binding site at K340 play a crucial role in modulating the impact of Sch9 on aflatoxin biosynthesis and stress response in A. flavus.

Added in the revised version: To investigate the potential function of Sch9 in the HOG and TOR pathways in A. flavus, we conducted a sensitivity analysis of the Δsch9 strain to rapamycin. Compared to the wild-type strain, the sch9 deletion strain exhibited a slightly higher sensitivity to rapamycin (Fig. 4B, 4F). The MAKP Hog1 is an element of the high osmolarity glycerol (HOG) pathway. To verify the connection between Sch9 and the HOG-MAPK pathway, the phosphorylation levels of Hog1 kinase were measured under osmotic stress (1M NaCl). Western blot analysis showed that the phosphorylation level of Hog1 increased significantly in the Δsch9 strain (Fig. 4E). This finding suggested that Sch9 plays a pivotal role in modulating the response of Hog1 to osmotic stress. Taken together, these results reflected that Sch9 regulates osmotic stress response via the HOG pathway in A. flavus.

Added in the revised version: TapA and TipA participate in regulation of sclerotia development and cell wall stress response

Added in the revised version: Tap42, a protein that associates with phosphatase 2A, serves as the direct target of the Tor kinase in yeast. Tap42 plays a crucial role in various Tor functions, particularly in the regulation of transcriptional processes [47]. The genome of A. flavus contains a homolog of tap42 named tapA (AFLA_092770), which is predicted to encode a 355 amino-acid protein (sharing 31.9% identity with S. cerevisiae Tap42). To further determine the function of the TapA regulator in the TOR signaling pathway of A. flavus, we conducted targeted gene deletion experiments on tapA. We failed to obtain a null mutant, therefore we generated the tapA+/− strain, which contains one copy of the tapA gene (Table S2). Compared to the wild-type strain, the tapA+/− strain exhibited diminished hyphal growth and conidiation (Fig. 5A, 5C), indicating that TapA is important for vegetative growth and conidiation. The results of the rapamycin sensitivity test indicated that the tapA+/− strain displayed reduced sensitivity to rapamycin (Fig.5A, 5E). In addition, the TapA+/− strain was unable to generate any sclerotia (Fig. 5B, 5F), and almost lost its capability in AFB1 production compared to the wild-type strain (Fig. 5D, 5G), these findings indicated that TapA is essential for sclerotia formation and biosynthesis of AFB1. The tapA+/− mutant was more insensitive to cell wall stress, including CFW and CR (Fig. S8). Furthermore, the levels of nsdC, nsdD, and sclR expression were significantly decreased in tapA+/− strain compared to those in the wild-type strain (Fig. S9). The transcriptional level of the tapA gene in the tapA+/− strain was drastically downregulated compared to the wild-type strain (Fig. S9). These findings have demonstrated the significantly positive impact of TapA on the growth, sporulation, sclerotia formation, cell wall stress, and AF biosynthesis in A. flavus.

Removed in the revised version: (F) Relative transcript levels of sclerotia-related genes in WT and ΔTapA1+/TapA2−strains.

Removed in the revised version: (G) Relative quantification of AFB1 production as in (C).

Removed in the revised version: (H) Relative transcript levels of aflatoxin biosynthesis regulatory and structural genes in WT and ΔTapA1+/TapA2−strains.

Removed in the revised version: TipA regulates growth, conidia, sclerotia and aflatoxin biosynthesis

Earlier version: In S. cerevisiae, Tip41 was identified as a Tap42-interacting protein and has been measured to function as a negative regulator of Tap42, which downregulates TORC1 signaling via activation of PP2A phosphatase [35]. The putative Tip41 orthologue in A. flavus named TipA(AFLA_047310). The amino acid sequence of TipA shares 40.49% similarity to S. cerevisiae Tap41. To elucidate the function of TOR signaling pathway protein TipA in A. flavus, we generated the TipA deletion mutant. The ΔTipA strain exhibited smaller diameter colonies on PDA compared to the WT strain, and the number of conidia produced by ΔTipA strain was significantly reduced than that by the WT strain on PDA (Fig. 5A). In addition, the ΔTipA strain was unable to produce any sclerotia (Fig. 5B), and TLC analyses indicated that the deletion of TipA led to a considerable decrease in AFB1 production on YES (Fig. 5C). Meanwhile, the ΔTipA strain exhibited decreased sensitivity to rapamycin (Fig. 5A). Moreover, the expression levels of genes nsdC and nsdD showed a decrease in ΔTipA strain compared to WT (Fig. 5F). These results indicated that TipA is important for growth, conidia, sclerotial and aflatoxin biosynthesis in A. flavus.

Revised version: In S. cerevisiae, Tip41 was identified as a Tap42-interacting protein and has been measured to function as a negative regulator of Tap42, which downregulates TORC1 signaling via activation of PP2A phosphatase [48]. The putative orthologue of tip41 in A. flavus named tipA (AFLA_047310). The amino acid sequence of TipA exhibits a similarity of 40.49% to Tap41 in S. cerevisiae. To elucidate the function of the TOR signaling pathway protein TipA in A. flavus, we generated the tipA deletion mutants. The ΔtipA strain exhibited a slightly reduced number of conidia produced on PDA compared to the wild-type strain (Fig. 5A). Meanwhile, the ΔtipA strain exhibited decreased sensitivity to rapamycin (Fig. 5A). The ΔtipA strain was unable to produce any sclerotia (Fig. 5B). In addition, the ΔtipA strain was more sensitive to cell wall and cell membrane stress, including CR and SDS (Fig. S8). Additionally, the expression levels of nsdC and nsdD exhibited a decrease in ΔtipA strain compared to the wild-type strain (Fig. S9). These results indicated that TipA is crucial for the sclerotial formation and cell wall stress in A. flavus.

Earlier version: Phosphatase SitA and Ppg1 are involved in growth, conidiation and sclerotia formation

Revised version: Phosphatase SitA and Ppg1 are involved in growth, conidiation, and sclerotial formation

Earlier version: In S. cerevisiae, Tap42 also interacts with the catalytic subunits of Type 2A phosphatases, such as Pph21, Pph22, Pph3, Sit4 and Ppg1, which is controlled by the Tor signaling pathway [36]. Phosphatase SitA(AFLA_108450) and Ppg1(AFLA_132610) were identified in A. flavus, which encode putative orthologs of Sit4 and Ppg1 in S. cerevisiae. To determine the biological functions of phosphatase SitA and Ppg1 in A. flavus, we disrupted the SitA and Ppg1 by homologous recombination strategy. To further study the relationship between phosphatase SitA and Ppg1, we also constructed complement and double deletion mutants. The vegetative development phenotype analysis showed that the ΔSitA, ΔPpg1 and ΔSitA/ΔPpg1 strains grew significantly more slowly than the WT strain on PDA (Fig. 6A). In addition, the ΔSitA, ΔPpg1 and ΔSitA/ΔPpg1 strains formed lower number of conidiophores (Fig. 6A). To gain further insight into the role of SitA and Ppg1 in conidiation, qRT-PCR was used to determine the transcript levels of conidial synthesis-related genes (abaA and brlA). The expression levels of the abaA and brlA genes were significantly downregulated in the ΔSitA, ΔPpg1 and ΔSitA/ΔPpg1 mutants compared to WT and complemented strains(Fig. 6G). The sclerotial formation analysis results showed that the ΔSitA, ΔPpg1 and ΔSitA/ΔPpg1 mutants were unable to produce any sclerotia (Fig. 6B).To further prove these findings, we conducted qRT-PCR to determine the expression levels of three regulators (NsdC, NsdD and SclR) of sclerotia development. Our results showed that the expression levels of NsdC, NsdD and SclR were all downregulated in the ΔSitA, ΔPpg1 and ΔSitA/ΔPpg1 mutants(Fig. 6H), which is consistent with the sclerotia production observed in the ΔSitA, ΔPpg1 and ΔSitA/ΔPpg1. These results showed that SitA and Ppg1 are crucial for hyphal development, conidiation and sclerotia formation in A. flavus.

Revised version: In S. cerevisiae, Tap42 has been identified to interact with the catalytic subunit of type 2A protein phosphatases, including Pph3, Pph21, and Pph22, as well as the type 2A-like phosphatases Sit4 and Ppg1[19]. Phosphatase sitA (AFLA_108450) and ppg1 (AFLA_132610) were identified in A. flavus, encoding putative orthologs of sit4 and ppg1 in S. cerevisiae. To determine the biological functions of phosphatase SitA and Ppg1 in A. flavus, we disrupted the sitA and ppg1 by homologous recombination strategy. To further investigate the relationship between phosphatase SitA and Ppg1, we additionally generated complemented and double deletion mutants. The vegetative development phenotype analysis showed that the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains grew significantly more slowly than the wild-type strain on PDA (Fig. 6A, 6D). Additionally, the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains were more sensitive to rapamycin (Fig. 6A, 6E). In addition, the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains formed lower number of conidiophores (Fig. 6A). To gain further insight into the role of SitA and Ppg1 in the process of conidiation, qRT-PCR was employed to determine the transcript levels of conidial synthesis-related genes (abaA and brlA). The expression levels of the abaA and brlA genes were significantly downregulated in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants when compared to the wild-type and complemented strains (Fig. 6G). The analysis of sclerotial formation revealed that the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants exhibited a complete inability to produce any sclerotia (Fig. 6B). To further prove these findings, we conducted qRT-PCR to determine the expression levels of three regulators (nsdC, nsdD, and nclR) of sclerotia development. Our findings indicated that the expression levels of nsdC, nsdD, and sclR were all significantly reduced in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants (Fig. 6H), which is consistent with the observed decrease in sclerotia production in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants. These results showed that SitA and Ppg1 are crucial for hyphal development, conidiation, and sclerotia formation in A. flavus.

Earlier version: Phosphatases SitA and Ppg1 are involved in growth, conidiation and aflatoxin biosynthesis in A. flavus.

Revised version: The impact of the phosphatases SitA and Ppg1 on the growth, conidiation, and aflatoxin biosynthesis in A. flavus

Earlier version: (A) The colony morphology of A. flavus WT, single knockout and double knockout strains of SitA and Ppg1 cultured on PDA medium amended with 100 ng/mL rapamycin at 37℃ for 5 days.

Revised version: (A) Colony morphology of the WT, single knockout, and double knockout strains cultured on PDA medium amended with 100 ng/mL rapamycin at 37℃ for 5 days.

Earlier version: (B) colony morphology of A. flavus WT, single knockout strain and double knockout strain of SitA and Ppg1 cultured on WKM medium at 37℃ for 7 days.

Revised version: (B) Colony morphology of the WT, single knockout, and double knockout strains cultured on WKM medium at 37℃ for 7 days.

Earlier version: (C) TLC detection of WT, single knockout and double knockout strains of SitA and Ppg1 cultured in YES liquid medium at 29℃ for 6 days.

Revised version: (C) TLC analysis of AFB1 production from the WT, single knockout, and double knockout strains cultured in YES liquid medium at 29℃ for 6 days.

Earlier version: (D) Comparison of conidial production among WT, single knockout and double knockout strains on PDA media.

Revised version: (D) Growth diameter of the WT, single knockout, and double knockout strains on PDA media.

Earlier version: (E) Sclerotia statistics of WT, single knockout and double knockout strains.

Revised version: (E) The growth inhibition rate of the WT, single knockout, and double knockout strains under rapamycin stress.

Earlier version: (F) Quantitative analysis of AFB1 of WT, single knockout and double knockout strains.

Revised version: (F) AFB1 quantitative analysis of the WT, single knockout, and double knockout strains.

Earlier version: (G) Relative transcript levels of conidia synthesis genes in WT, single knockout and double knockout strains.

Revised version: (G) Relative expression levels of conidia synthesis genes in the WT, single knockout, and double knockout strains.

Earlier version: (H) Relative transcript levels of sclerotia synthesis genes in WT, single knockout and double knockout strains.

Revised version: (H) Relative expression levels of sclerotia synthesis genes in the WT, single knockout, and double knockout strains.

Earlier version: (I) Relative transcript levels of aflatoxin biosynthesis genes in WT, single knockout and double knockout strains.

Revised version: (I) Relative expression levels of aflatoxin biosynthesis genes in the WT, single knockout, and double knockout strains.

Earlier version: SitA and Ppg1 play different roles in regulating aflatoxin biosynthesis

Revised version: SitA and Ppg1 play crucial roles in aflatoxin biosynthesis and pathogenicity

Earlier version: Previous studies have demonstrated that protein phosphorylation is critical for the regulation of AFs biosynthesis, which are the toxic and carcinogenic secondary metabolites of A. flavus [37]. We assayed AFB1 biosynthesis in the ΔSitA, ΔPpg1 and ΔSitA/ΔPpg1 mutants, and TLC analyses indicated that ΔSitA and ΔSitA/ΔPpg1 strains produced less AFB1, while the level of AFB1 produced by ΔPpg1 was markedly increased compared to that of the WT and complemented strains (Fig. 6C). The expression levels of aflatoxin synthesis-related genes (structural genes: aflO, aflP, aflQ and aflJ, and regulatory genes: aflR and aflS) were determined by qRT-PCR in the ΔSitA, ΔPpg1 and ΔSitA/ΔPpg1 mutants, and the result showed that the expression levels of those genes were drastically downregulated in the ΔSitA and ΔSitA/ΔPpg1 mutants, while upregulated in ΔPpg1 mutants when compared to those of the WT and complemented strains(Fig. 6I). These results indicated that SitA and Ppg1 play different roles in regulating aflatoxin biosynthesis in A. flavus.

Revised version: We assayed AFB1 biosynthesis in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants. TLC analyses revealed that the ΔsitA and ΔsitA/Δppg1 strains exhibited a decrease in AFB1 production, whereas the Δppg1 strain showed a significant increase in AFB1 production compared to the wild-type and complemented strains (Fig. 6C, 6F). The expression levels of genes related to aflatoxin synthesis, including structural genes (aflO and aflP) and regulatory genes (aflR and aflS), were determined by qRT-PCR in the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants. The results revealed a significant downregulation of these genes in the ΔsitA and ΔsitA/Δppg1 mutants, while upregulated in the Δppg1 mutants, compared to the wild-type and complemented strains (Fig. 6I). These results indicated that SitA and Ppg1 play different roles in regulating aflatoxin biosynthesis in A. flavus.

Removed in the revised version: SitA and Ppg1 are important for pathogenicity

Earlier version: To investigate the pathogenicity of phosphatases SitA and Ppg1 on crops, we conducted pathogenicity tests on peanut and maize. The results showed that WT and complement strains resulted in full virulence on all peanut and maize seeds, while the ΔSitA, ΔPpg1 and ΔSitA/ΔPpg1 mutants exhibited more flawed colonization on peanut and maize seeds (Fig. 7A). In addition, the number of conidial productions of these mutants on the infected peanut seeds also dramatically decreased compared to those of WT and complement strains (Fig. 7C). TLC analyses indicated that the ΔPpg1 mutant produced more AFB1 on peanut and maize, while the level of aflatoxin produced by ΔSitA and ΔSitA/ΔPpg1 mutants on peanut and maize were significantly reduced compared to WT and complement strains (Fig. 7B and 7D). This is consistent with the above results of aflatoxin biosynthesis of all mutants in the medium. All these data showed that SitA and Ppg1 are important for crop seeds pathogenicity.

Revised version: We conducted pathogenicity tests on peanuts and maize to examine the pathogenicity of the phosphatases SitA and Ppg1 on crops. The results showed that the wild-type and complemented strains demonstrated complete virulence on all peanut and maize seeds. In contrast, the ΔsitA, Δppg1, and ΔsitA/Δppg1 mutants displayed compromised colonization on peanut and maize seeds (Fig. 7A). In addition, the number of conidial productions in these mutants on the infected peanut seeds also dramatically decreased compared to those of the wild-type and complemented strains (Fig. 7C). The thin-layer chromatography analyses revealed that the Δppg1 mutant exhibited increased production of AFB1 in peanuts and maize. Conversely, the levels of aflatoxin produced by the ΔsitA and ΔsitA/Δppg1 mutants on peanut and maize were significantly reduced compared to the the wild-type and complemented strains (Fig. 7B, 7D). This observation aligns with the above results of aflatoxin biosynthesis from all mutant strains in the YES medium. All these data showed that SitA and Ppg1 are crucial for crop seeds’ pathogenicity.

Earlier version: Phosphatase SitA and Ppg1 are important for pathogenicity in A. flavus.

Revised version: Pathogenicity analysis of the phosphatase SitA and Ppg1 in A. flavus

Earlier version: (A) Phenotypes of peanut and maize infected by A. flavus WT, single and double knockout strains of SitA and Ppg1.

Revised version: (A) Phenotypes of peanut and maize infected by the WT, single knockout, and double knockout strains.

Earlier version: (B) TLC detection of toxin from peanut and maize infected by WT, single knockout and double knockout strains.

Revised version: (B) TLC was used to detect the AFB1 from peanut and maize infected by the WT, single knockout, and double knockout strains.

Removed in the revised version: (C) Spores statistics from peanut and maize infected by WT, single and double knockout strains.

Earlier version: (D) Quantitative analysis of AFB1 from peanut and maize infected by WT, single knockout and double knockout st

Revised version: (C) Quantitative analysis of conidia production from peanut and maize infected by the WT, single knockout, and double knockout strains.

Added in the revised version: (D) Quantitative analysis of AFB1 from peanut and maize infected by the WT, single knockout, and double knockout strains.

SitA and Ppg1 are involved in regulating cell wall integrity

Earlier version: To explore the role of phosphatases SitA and Ppg1 in response to cell wall stress, all strains were cultured on PDA amended with CFW and CR. We found that the relative growth inhibition of ΔSitA, ΔPpg1 and ΔSitA/ΔPpg1 mutants induced by cell wall stress were significantly increased compared to WT and complement strains (Fig. 8A and 8B). Further studies by qRT-PCR showed that expression levels of genes associated with synthetic cell walls, such as chitin synthase genes chsA, chsB, chsC and chsD, were significantly decreased in the ΔSitA, ΔPpg1 and ΔSitA/ΔPpg1 strains (Fig. 8C). Western blot analysis showed that the phosphorylation level of Slt2 was higher in ΔSitA, ΔPpg1 and ΔSitA/ΔPpg1 than that in the WT strain (Fig. 8D). To sum up, the phosphatases SitA and Ppg1 maintained the integrity of the cell wall by positively regulating the expression of those genes.

Revised version: To explore the role of phosphatases SitA and Ppg1 in response to cell wall stress, all strains were cultivated on PDA amended with CFW, CR, and SDS. We found that the relative growth inhibition of ΔsitA, Δppg1 and ΔsitA/Δppg1 mutants induced by cell wall stress were significantly increased compared to the wild-type and complemented strains (Fig. 8A, 8B). Further investigations using qRT-PCR revealed a notable reduction in the expression levels of genes related to the synthesis of cell walls, specifically chitin synthase genes chsA, chsD, chsE, and β-1,6, in the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains (Fig. 8C). Western blot analysis also revealed that the phosphorylation level of Slt2 was significantly elevated in the ΔsitA, Δppg1, and ΔsitA/Δppg1 strains under cell wall stress (200 μg/ml CFW) (Fig. 8D). In conclusion, the phosphatases SitA and Ppg1 play a crucial role in regulating cell wall integrity.

Added in the revised version: Sensitivity of phosphatase SitA and Ppg1 to cell wall damaging agents in A. flavus

Removed in the revised version: Phosphatase SitA and Ppg1 are involved in regulating cell wall integrity.

Earlier version: (A) Morphology of A. flavus WT, single knockout and double knockout strains of SitA and Ppg1 grown on PDA media supplemented with 200 µg/mL CR, 200 µg/mL CFW, or 100 µg/mL SDS at 37°C for 5 days.

Revised version: (A) Morphology of the WT, single knockout, and double knockout strains grown on YES media supplemented with 200 µg/mL CR, 200 µg/mL CFW, or 100 µg/mL SDS at 37°C for 5 days.

(B) The growth inhibition rate of the WT, single knockout, and double knockout strains under cell wall and cell membrane stress.

Earlier version: (C) Expression levels of cell wall synthesis genes (chsA, chsD, chsE and β-1,6) in WT, single knockout and double knockout strains.

Revised version: (C) Relative expression levels of cell wall synthesis genes (chsA, chsD, chsE and β-1,6) in the WT, single knockout, and double knockout strains.

Earlier version: (D) Phosphorylation level of Slt2 in WT, single knockout and double knockout strains was detected by Western blotting.

Revised version: (D) The phosphorylation level of Slt2 in the WT, single knockout, and double knockout strains was detected with or without cell wall stress.

Earlier version: The phosphatase complex Nem1/Spo7 are important for growth, conidiation, aflatoxin and LD biogenesis

Revised version: The phosphatase complex Nem1/Spo7 are important for growth, aflatoxin biosynthesis, and LD biogenesis

Earlier version: Lipid droplets (LDs) are spherical organelles, which appear in the cytosol of some prokaryotic and most eukaryotic cells upon energy surplus. The FgTOR pathway regulates the phosphorylation status of FgNem1 mainly via the FgPpg1/Sit4-FgCak1/other kinase cascade in F. graminearum [7]. In this study, we observed that rapamycin-treated hyphae led to the accumulation of LD biogenesis(Fig. 9A), suggesting that the TOR pathway may regulate the LD biogenesis. However, the functions of LDs and the underlying regulatory mechanism of LD biosynthesis are still largely unknown. In order to explore the regulatory mechanism of the TOR pathway that regulates LD biogenesis in A. flavus, we examined LD accumulation in ΔSch9, ΔSitA and ΔPpg1. BODIPY staining assays showed that LD biogenesis was not induced in the ΔSitA and ΔPpg1 strains treated with rapamycin compared to those in the WT and ΔSch9 strains (Fig. 9A-D). These results indicated that LD biogenesis induced by rapamycin is largely dependent on the phosphatases SitA and Ppg1 in A. flavus. In F. graminearum, the phosphatase FgNem1/FgSpo7 complex was involved in induction of LD biogenesis by rapamycin. Phosphatase Nem1(AFLA_002649) and Spo7 (AFLA_028590) were identified in A. flavus, which encode putative orthologs of FgNem1 and FgSpo7 in F. graminearum. To investigate the function of Nem1 and Spo7 gene in A. flavus, ΔNem1 and ΔSpo7 mutants was constructed with homologous recombination. The ΔNem1 and ΔSpo7 mutants exhibited a significant decrease in vegetative growth and conidiation compared to WT(Fig. 10A). TLC assay and quantitative analysis showed a significantly decreased aflatoxin production in ΔNem1 and ΔSpo7 compared to WT(Fig. 10B). In addition, BODIPY staining showed no visible lipid droplets in the hyphae of ΔNem1 and ΔSpo7 after rapamycin treatment(Fig. 10C). Taken together, these results indicated that Nem1/Spo7 are involved in vegetative growth, conidiation, aflatoxin and LD biogenesis.

Revised version: Lipid droplets (LDs) are highly dynamic spherical organelles, which appear in the cytosol of most eukaryotic cell types[49]. Rapamycin treatment led to the accumulation of lipid droplets in various filamentous fungi [24]. The FgTOR pathway regulates the phosphorylation status of FgNem1 mainly via the FgPpg1/Sit4-FgCak1/other kinase cascade in F. graminearum [24]. In this study, we observed that the treatment of hyphae with rapamycin resulted in the accumulation of LD biogenesis (Fig. 9A). This finding suggested that the TOR pathway may regulate the LD biogenesis. However, the underlying regulatory mechanism of LD biosynthesis remain largely unknown in A. flavus. To investigate the regulatory mechanism of the TOR pathway in A. flavus that controls LD biogenesis, we examined LD accumulation in Δsch9, ΔsitA, and Δppg1 mutants. BODIPY staining assays revealed that LD biogenesis was not observed in the ΔsitA and Δppg1 strains upon treatment with rapamycin, in contrast to the wild-type and Δsch9 strains (Fig. 9A, 9B, 9C, 9D). These results indicated that LD biogenesis induced by rapamycin is predominantly reliant on the phosphatases SitA and Ppg1 in A. flavus. In F. graminearum, the phosphatase FgNem1/FgSpo7 complex was involved in the induction of LD biogenesis by rapamycin. Phosphatase nem1 (AFLA_002649) and spo7 (AFLA_028590) were identified in A. flavus, which encode putative orthologs of nem1 and spo7 in F. graminearum. To examine the role of the nem1 and spo7 genes in A. flavus, we generated Δnem1 and Δspo7 mutants using homologous recombination. The Δnem1 and Δspo7 mutants displayed a notable reduction in both vegetative growth and conidiation compared to the wild-type strain (Fig. 10A, 10C, 10D). TLC assay revealed a significantly decreased aflatoxin production in the Δnem1 and Δspo7 compared to the wild-type strain (Fig. 10B). In addition, BODIPY staining showed no visible lipid droplets in the hyphae of Δnem1 and Δspo7 after rapamycin treatment (Fig. 10E). Taken together, these results indicated that Nem1/Spo7 are involved in vegetative growth, conidiation, aflatoxin production, and LD biogenesis.

Phosphatase SitA and Ppg1 are involved in regulating lipid droplet biogenesis

Earlier version: (A) The intracellular lipid droplets were stained by boron dipyrromethene difluoride (BODIPY) in hyphae of the A. flavus WT strain and observed under a fluorescence microscopy. The Phenotype of lipid droplets accumulation in the mycelia of A. flavus WT strain treated with or without 100 ng/mL rapamycin cultured for 6 h. Bar, 10 µm.

Revised version: (A) The intracellular lipid droplets were stained by boron dipyrromethene difluoride (BODIPY) in hyphae of the WT strain and observed under fluorescence microscopy. The phenotype of lipid droplets accumulation in the mycelia of the WT strain treated with or without 100 ng/mL rapamycin for 6 h (Bar, 10 µm).

Earlier version: (B) The phenotype of lipid droplets accumulation in the mycelia of ΔSch9 strain treated with or without 100 ng/mL rapamycin cultured for 6 h.Bar, 10 µm.

Revised version: (B) Phenotype of lipid droplets accumulation in the mycelia of the Δsch9 strain treated with or without 100 ng/mL rapamycin for 6 h (Bar, 10 µm).

Earlier version: (C) The phenotype of lipid droplets accumulation in the mycelia of ΔSitA strain treated with or without 100 ng/mL rapamycin cultured for 6 h.Bar, 10 µm.

Revised version: (C) Phenotype of lipid droplets accumulation in the mycelia of the ΔsitA strain treated with or without 100 ng/mL rapamycin for 6 h (Bar, 10 µm).

Earlier version: (D) The phenotype of lipid droplets accumulation in the mycelia of ΔPpg1 strain treated with or without 100 ng/mL rapamycin cultured for 6 h. Bar, 10 µm.

Revised version: (D) Phenotype of lipid droplets accumulation in the mycelia of the Δppg1 strain treated with or without 100 ng/mL rapamycin for 6 h (Bar, 10 µm).

Earlier version: Phosphatase complex Nem1/Spo7 play an important role in growth, conidiation and aflatoxin and lipid droplet biogenesis.

Revised version: Phosphatase complex Nem1/Spo7 play significant role in growth, conidiation, aflatoxin, and lipid droplet biogenesis

Earlier version: (A) Phenotype of the A. flavus WT, ΔNem1 and ΔSpo7 strains grown on PDA amended with 100 ng/mL rapamycin at 37°C for 5 days.

Revised version: (A) Phenotype of the WT, Δnem1, and Δspo7 strains grown on PDA amended with 100 ng/mL rapamycin at 37°C for 5 days.

Earlier version: (B) TLC assay of AFB1 production by the A. flavus WT, ΔNem1 and ΔSpo7 strains cultured in YES liquid medium at 29°C for 6 days.

Revised version: (B) TLC assay of AFB1 production from the WT, Δnem1, and Δspo7 strains cultured in YES liquid medium at 29°C for 6 days.

Earlier version: (C) Statistical analysis of the diameter and sporulation by the A. flavus WT, ΔNem1 and ΔSpo7 strains.

Revised version: (C) Growth diameter of the WT, Δnem1, and Δspo7 strains on PDA media.

Added in the revised version: (D) Statistical analysis of the sporulation by the WT, Δnem1, and Δspo7 strains.

Earlier version: (D)The Phenotype of lipid droplets accumulation in the mycelia of the A. flavus WT, ΔNem1 and ΔSpo7 strains treated with or without 100 ng/mL rapamycin cultured for 6 h. Bar, 10 µm.

Revised version: (E) Phenotype of lipid droplets accumulation in the mycelia of the WT, Δnem1, and Δspo7 strains treated with or without 100 ng/mL rapamycin for 6 h (Bar, 10 µm).

Discussion

Earlier version: Aflatoxins (AFs) are a class of toxic secondary metabolites produced by A. flavus, including aflatoxin B1, B2, G1 and G2, among which AFB1 has been regarded as the most pathogenic and toxic aflatoxin [38]. Therefore, understanding the regulatory mechanism of fungal secondary metabolite biosynthesis would help in exploiting aflatoxin control in A. flavus. In this study, we demonstrated that the target of the rapamycin (TOR) signaling pathway plays critical roles in regulating growth, sporulation, sclerotia and aflatoxin production in A. flavus. Through functional studies of nine genes in the TOR signaling pathway, we explored the conservation and complexity of the TOR signaling pathway in A. flavus. In addition, we constructed a crosstalk network between TOR and other signaling pathways (MAPK, CWI) and analyzed the specific regulatory process. We proposed that the TOR signaling pathway interacts with other signaling pathways (MAPK, CWI, calcineurin-CrzA pathway) to regulate the responses to various environmental stresses(Fig. 11). Our results might provide new insights into the biological control of A. flavus.

Revised version: Aflatoxins (AFs) are a class of toxic secondary metabolites synthesized by A. flavus, including aflatoxin B1, B2, G1, and G2, among which AFB1 has been regarded as the most pathogenic and toxic aflatoxin[50]. Therefore, understanding the regulatory mechanisms involved in fungal secondary metabolites biosynthesis would be instrumental for effectively addressing aflatoxin control in A. flavus. In this study, we demonstrated that the target of rapamycin (TOR) signaling pathway plays a pivotal role in regulating growth, sporulation, sclerotia formation, aflatoxin production, and various stress response in A. flavus. Through conducting functional studies on genes within the TOR signaling pathway, we explored the conservation and complexity exhibited by the TOR signaling pathway in A. flavus. In addition, we constructed a crosstalk network between the TOR and other signaling pathways (HOG, CWI) and analyzed the specific regulatory process. We proposed that the TOR signaling pathway interacts with multiple signaling pathways, including HOG and CWI pathway, to modulate the cellular responses to diverse environmental stresses (Fig. 11).

Earlier version: The proposed model of target of the rapamycin (TOR) pathway in A. flavus.

Revised version: The proposed model of the target of rapamycin (TOR) pathway in A. flavus

Earlier version: Based on the above results, we hypothesized the TOR signaling pathway model in A. flavus. Rapamycin forms a complex with FKBP3, and this complex can bind to the Tor kinase to inhibit its normal function. Sch9 as a downstream component of the Tor kinase, regulates sclerotia formation, aflatoxin biosynthesis, the TOR and HOG signaling pathways. In addition, We speculate that the Sch9 as a element for the calcineurin-CrzA pathway in response to calcium stress. As another important target of the Tor kinase, Tap42-phosphatase complex is involved in the regulation of growth, sporulation, sclerotia, aflatoxin production, CWI signaling pathways, lipid droplet synthesis and other processes. In conclusion, the TOR signaling pathway plays an important role in the vegetative development, stress responses, aflatoxin biosynthesis and pathogenicity in A. flavus.

Revised version: Based on the above results, we propose a hypothesis regarding the TOR signaling pathway model in A. flavus. Rapamycin forms a complex with Fkbp3, and this particular complex can bind to the Tor kinase, thereby impeding its regular functionality. Sch9, functioning as a downstream element of the Tor kinase, regulates aflatoxin biosynthesis and HOG signaling pathway. As another important target of the Tor kinase, TapA-phosphatase complex is involved in the regulation of growth, sporulation, sclerotia, aflatoxin production, CWI signaling pathway, lipid droplet synthesis, and other processes. In conclusion, the TOR signaling pathway plays a crucial role in various aspects of A. flavus, including vegetative development, stress response, aflatoxin biosynthesis, and pathogenicity.

Earlier version: We found that rapamycin not only affects the growth and sporulation of A. flavus, but also strongly inhibits sclerotia production and aflatoxin biosynthesis (Fig. 1A-1C). This inhibitory effect has been observed in a variety of fungi, including Candida albicans [39], Cryptococcus neoformans [39], Podospora anserine [40], Aspergillus nidulans [26], Magnaporthe oryzae [8], and Fusarium graminearum [21], suggesting that the function of the TOR signaling pathway may be conserved in filamentous fungi. The number of FK506-binding proteins varies among species: there are four FKBPs protein in S. cerevisiae and three FKBPs protein in S. pombe [34]. In S. cerevisiae, rapamycin can bind to FKBP protein, irreversibly inhibit the G1 phase of the cell cycle and control cell growth [46]. The entomopathogenic fungus Beauveria bassiana contains three putative FKBP genes, and disruption of BbFKBP12 significantly increases the resistance to rapamycin and FK506 [41]. The plant pathogenic fungus F. graminearum harbors three FKBPs protein, among which deletion of FgFKBP12 leads to resistance to rapamycin [21]. In this study, we showed that FKBP3 is involved in sclerotia and aflatoxin biosynthesis and rapamycin resistance in A. flavus(Fig. 2A-2C), and the conserved site K19 of FKBP3 regulated the aflatoxin biosynthesis and rapamycin resistance(Fig. 2E-2G). Previous research has shown that FKBP12 was characterized as the receptor of FK506 and rapamycin, which has been the most extensively documented in different fungi. The deletion of FKBP12 confers the resistance to rapamycin in C. albicans [39] and C. neoformans [39]. Similar phenotypes are observed in other fungi, such as F. graminearum[21], B. bassiana [41]and Mucor circinelloides[30], and the FKBP12 mutants have significantly increased resistance to rapamycin and FK506 in these filamentous fungi. In Botrytis cinerea [28], the deletion of BcFKBP12 cause a reduction of the virulence of strain T4, while knockout of the FKBP12 gene did not affect the pathogenic development of the strain B05.10. Thus, we proposed that these FKBPs may perform similar functions in fungi.

Revised version: We have observed that rapamycin has a significant impact on the growth and sporulation of A. flavus (Fig. 1A). Additionally, it exhibits a strong inhibitory effect on sclerotia production and aflatoxin biosynthesis (Fig. 1B, 1C). The inhibitory effect has been documented in various fungal species, including Candida albicans [51], Cryptococcus neoformans [51], Podospora anserine [27], A. nidulans [28], Magnaporthe oryzae [39], and F. graminearum [23]. This suggested that the TOR signaling pathway may have a conserved function in filamentous fungi. The number of FK506-binding proteins varies among different species: there are four Fkbps protein in S. cerevisiae and three Fkbps protein in S. pombe [52]. In S. cerevisiae, rapamycin can bind to Fkbp protein, resulting in the irreversible inhibition of the G1 phase of the cell cycle and the regulation of cell growth [53]. In A. nidulans, the fkbp12 gene homologue fprA displays considerable sequence similarity with homologues identified in the S. cerevisiae [28]. In this study, based on the homologous alignment of the protein sequence of FprA, it was observed that four fkbp genes present in A. flavus. Among these, deletion of fkbp3 exhibited enhanced tolerance to rapamycin (Fig. 2A). So we speculated that rapamycin forms a complex with Fkbp3 in A. flavus, subsequently targeting downstream target genes in order to exert its biological function. We found that Fkbp3 is involved in sclerotia and aflatoxin biosynthesis in A. flavus (Fig. 2A, S3), and subsequent analysis revealed that the underlying mechanism may be associated with succinylation modification. The succinylation sites of Fkbp3 were identified through the utilization of succinylation proteomics data. By constructing point mutant strains and conducting phenotype experiments, we found that the succinylation of Fkbp3 at residue K19 is involved in rapamycin resistance and aflatoxin biosynthesis (Fig. 2E, 2G). Previous studies have demonstrated that Fkbp12 has been identified as the receptor for FK506 and rapamycin, which inhibits calcineurin and target of rapamycin complex 1 (TORC1), respectively [53]. The deletion of Fkbp12 confers the resistance to rapamycin in C. albicans [51] and C. neoformans [51]. The entomopathogenic fungus Beauveria bassiana contains three putative fkbp genes, and disruption of Bbfkbp12 significantly increases the resistance to rapamycin and FK506 [54]. The plant pathogenic fungus F. graminearum harbors three Fkbps proteins, among which deletion of Fgfkbp12 leads to resistance to rapamycin [23]. Thus, we propose that these Fkbps may exhibit analogous functionalities in conferring resistance to rapamycin. In Botrytis cinerea, deletion of Bcfkbp12 cause a reduction of the virulence of strain T4, while knockout of the fkbp12 gene did not affect the pathogenic development of the strain B05.10 [32]. These findings demonstrated that the fkbp12 homologous genes exhibited conservation in terms of rapamycin resistance. However, there is significant variation in other biological functions among fungal species, which may be attributed to species or strain evolutionary divergence.

Earlier version: The Tor kinase acts as a core element of the TOR signaling pathway, which is evolutionarily conserved in fungi, plants and mammals [45]. TOR assembles into two distinct multiprotein complexes: TORC1 and TORC2. A defining feature of TORC1 is the interaction of TOR with KOG1 (Raptor in mammals) and its sensitivity to a rapamycin-FKBP12 complex [61]. Tor kinase is widely conserved both structurally and as the target of FKBP-rapamycin. Almost all other eukaryotes, including plants, worms, flies, and mammals, have a single Tor gene, whereas S. cerevisiae has two Tor genes [47]. Tor kinase is widely conserved both structurally and as the target of FKBP-rapamycin, which have similar domains including HEAT-like, FAT, FRB, kinase and FATC domains. The FRB domain is the binding region of the FKBP-Rapamycin complex, and rapamycin resistance is caused by the disruption of the FRB domain[47]. The homology comparison and phylogenetic analysis revealed that there is the only Tor kinase gene predicted in A. flavus. To determine its function, we generated the Tor single-copy deletion mutant by homologous recombination. We found that Tor was involved in the vegetative growth and asexual development of A. flavus(Fig. 3C). In addition, Tor responds to various stresses including osmotic and oxidative stresses, cell wall-damaging agents and rapamycin(Fig. 3F, S6). Our results revealed that Tor kinase and its FRB domain play very important roles in the morphogenesis, mycotoxin biosynthesis, and multiple stress responses of A. flavus. Previous findings demonstrated that Tor kinase could function as a molecular switch connecting an iron cue to defend against pathogen infection in C. elegans [57]. In A. fumigatus, Tor kinase represents a central regulatory node controlling genes and proteins involved in nutrient sensing, stress response, cell cycle progression, protein biosynthesis and degradation, as well as adaptation to low iron conditions[58]. In F. graminearum, the Tor point mutation (S1866L) prevents Fkbp12-rapamycin complex from binding to the FgFRB domain of FgTor [21]. We also constructed a point mutant strain TORS1904L, but the sensitivity to rapamycin is consistent with the WT strain, which may be related to the functional variances of the Tor site (S1866L) in different species.

Revised version: The Tor kinase functions as a fundamental component of the TOR signaling pathway, which exhibits evolutionary conservation across fungi, plants, and mammals [55]. TOR assembles into two distinct multi-protein complexes known as TORC1 and TORC2. A characteristic that distinguishes TORC1 is its interaction with KOG1 (Raptor in mammals) and its sensitivity to the rapamycin-FKBP12 complex [56]. Almost all other eukaryotes, including plants, worms, flies, and mammals, have a single tor gene, whereas S. cerevisiae has two tor genes [45]. The Tor kinase is widely conserved both in terms of its structural characteristics and its role as the target of Fkbp-rapamycin. Typically, the Tor kinase exhibits comparable structural characteristics, including the HEAT-like, FAT, FRB, kinase, and FATC domains. The FRB domain is the binding region of the Fkbp-rapamycin complex, and rapamycin resistance is usually caused by the disruption of the FRB domain [45]. The comparison of homology and the analysis of phylogenetics revealed that there is only one predicted tor gene in A. flavus. To ascertain its functionality, we generated the tor single-copy deletion mutant by homologous recombination. We discovered that the Tor kinase is involved in the vegetative growth and asexual development of A. flavus (Fig. 3B, 3C). In addition, the tor+/− mutant shows a high sensitivity to rapamycin (Fig. 3C). Our findings have demonstrated the significant involvement of Tor kinase and its FRB domain in the processes of morphogenesis, aflatoxin biosynthesis, and multiple stress responses in A. flavus. We speculate the common expectation that the deletion of the tor and FRB domain should result in insensitivity and resistance to rapamycin, as it disrupts the binding site for Fkbp-rapamycin. However, we observed that the tor+/− and FRB domain-deleted mutant were more sensitive to rapamycin. This intriguing result suggests that there are additional factors or complexities involved in TOR signaling pathway regulation in A. flavus. We hypothesize that this result is related to the double copy of the tor gene. Previous findings demonstrated that the Tor kinase could function as a molecular switch connecting an iron cue to defend against pathogen infection in C. elegans [57]. In A. fumigatus, the Tor kinase represents a central regulatory node controlling genes and proteins involved in nutrient sensing, stress response, cell cycle progression, protein biosynthesis, and degradation, as well as adaptation to low iron conditions[58]. In summary, the Tor kinase acts as a central regulator in the TOR signaling pathway, playing a crucial role in the regulation of the cell cycle, protein synthesis, and cellular energy metabolism.

Earlier version: In S. cerevisiae, the protein kinase Sch9 is a major downstream effector of TORC1, which plays important roles in stress resistance, longevity and nutrient sensing [59]. In addition, the TORC1-Sch9 pathway acts as a crucial mediator of chronological lifespan [50]. In filamentous fungi, there is a complex network of interactions between TOR and HOG pathways. In F. graminearum [7], FgSch9 and FgHog1 mutants exhibited increased sensitivity to osmotic and oxidative stresses, and this defect was more severe in the FgSch9/FgHog1 double mutant. The phosphorylation level of FgHog1 increased significantly in ΔFgSch9, and the affinity capture-MS assay showed that TOR was one of the proteins co-purifying with FgSch9. These results suggest that FgSch9 serves as a mediator for the TOR and HOG pathways[7]. In addition, FgSch9 regulates hyphal differentiation, asexual development, virulence and DON biosynthesis, which is also involved in the regulation of sensitivity to rapamycin, cell wall-damaging agents, fungicide, osmotic and oxidative stresses[7]. In A. fumigatus, the SchA mutant was sensitive to rapamycin, high concentrations of calcium and hyperosmotic stress, and SchA was involved in iron metabolism. The SchA mutant showed increased phosphorylation of SakA upon osmotic stress[51]. In this study, we discovered that Sch9 is involved in regulating aflatoxin biosynthesis (Fig.4C), and responses to various stresses including osmotic stress, cell wall-damaging agents and rapamycin (Fig.4B). The S_TKc domain deletion strain (Sch9ΔS_TKc), S_TK_X domain deletion strain (Sch9ΔS_TK_X) and Sch9K340A mutant exhibited similar phenotype with ΔSch9 strain(Fig. 4C-4D). These findings revealed that Sch9 plays an important role in the biological synthesis of aflatoxin and multiple stresses mainly via its kinase domain, kinase-C domain and ATP binding site k340. We also found that the phosphorylation level of Hog1 increased significantly in the ΔSch9 strain(Fig. 4E). Based on the above results, we speculated that Sch9 serves as a mediator for the complex network of interactions amongst the TOR and HOG signaling pathways.

Revised version: In S. cerevisiae, the protein kinase Sch9 is a major downstream effector of TORC1, and this kinase is known to have crucial functions in stress resistance, longevity, and nutrient sensing [46]. In addition, the TORC1-Sch9 pathway acts as a crucial mediator of chronological lifespan [59]. In filamentous fungi, an intricate network of interactions exists between the TOR and HOG pathways. In F. graminearum [38], FgSch9 and FgHog1 mutants displayed heightened susceptibility to osmotic and oxidative stresses. The phosphorylation level of FgHog1 increased significantly in FgSch9 mutant. Additionally, the affinity capture-MS assay showed that the Tor kinase was among the proteins co-purifying with FgSch9. These results suggest that FgSch9 serves as an intermediary for the TOR and HOG signaling pathways [38]. Furthermore, FgSch9 plays a crucial role in the regulation of various biological processes such as hyphal differentiation, asexual development, virulence, and DON biosynthesis[38]. In A. fumigatus, the SchA mutant was sensitive to rapamycin, high concentrations of calcium and hyperosmotic stress, and SchA was involved in iron metabolism. The SchA mutant exhibited enhanced phosphorylation of SakA in response to osmotic stress [51]. In this study, we discovered that Sch9 is involved in regulating aflatoxin biosynthesis (Fig.4C), and responses to osmotic stress and rapamycin stress (Fig.4B, 4D). The S_TKc domain deletion strain (sch9ΔS_TKc), S_TK_X domain deletion strain (sch9ΔS_TK_X), and sch9K340A mutant displayed a similar phenotype to the Δsch9 strain (Fig. 4B, 4C, 4D). These findings revealed that Sch9 plays a significant role in the biological synthesis of aflatoxin and in responding to various stresses. This role is primarily mediated through its kinase domain, kinase-C domain, and ATP binding site k340. We also found that the phosphorylation level of Hog1 increased significantly in the Δsch9 strain under osmotic stress (Fig. 4E). Based on the above results, we hypothesized that Sch9 was involved in the osmotic stress response by modulating the HOG pathway.

Earlier version: In S. pombe, Tip41 (a Tap42-interacting protein) plays an important role in the cellular responses to nitrogen sources by regulating type 2A phosphatases [42]. In S. cerevisiae, Tap42-phosphatase complexes associate with TORC1, and Tip41 downregulates TORC1 signaling via activation of PP2A phosphatase [34]. Unlike yeast Tip41, TIPRL, a mammalian Tip41-like protein, has a positive effect on mTORC1 signaling through the interaction with PP2Ac [43]. While the interaction between Tip41 and Tap42 was not observed in the filamentous fungi. The Y2H assay showed that the FgTip41 interacts with phosphatase FgPpg1 rather than FgTap42, while the FgTap42 bind with phosphatases FgPp2A, implying that FgTap42, FgPpg1 and FgTip41 form a heterotrimer in F. graminearum [21]. FgTip41 is associated with mycotoxin production and pathogenicity in F. graminearum, but ΔFgTip41 had no detectable changes in sensitivity to rapamycin[21]. In the rice blast fungus M. oryzae, MoTip41 does not interact with the MoTap42, but MoTip41 can bind with phosphatase MoPpe1 to mediate crosstalk between TOR and cell wall integrity signaling[9]. In this study, we found that the TapA and TipA were essential for growth, conidia, sclerotia and aflatoxin biosynthesis in A. flavus(Fig. 5A-5C). The TapA single-copy deletion mutant exhibited decreased sensitivity to rapamycin. However, the TipA deletion mutant had no detectable changes in sensitivity to rapamycin. These results indicated that the functions of TapA and TipA in A. flavus were similar to its ortholog in the filamentous fungi.

Revised version: In S. pombe, the protein Tip41, which interacts with Tap42, plays a significant role in the cellular responses to nitrogen sources by regulating type 2A phosphatases [60]. In S. cerevisiae, Tap42-phosphatase complexes associate with TORC1, whereas Tip41 functions to attenuate TORC1 signaling by stimulating the PP2A phosphatase [48]. Unlike its yeast counterpart Tip41, TIPRL, a mammalian Tip41-like protein, exerts a positive influence on mTORC1 signaling by interacting with PP2Ac [61]. While the interaction between Tip41 and Tap42 was not observed in the filamentous fungi. The yeast two-hybrid (Y2H) assay showed that FgTip41 interacts with the phosphatase FgPpg1 rather than FgTap42, while the FgTap42 bind with phosphatases FgPp2A, implying that FgTap42, FgPpg1, and FgTip41 form a heterotrimer in F. graminearum [23]. In the rice blast fungus M. oryzae, MoTip41 does not interact with the MoTap42. However, MoTip41 has the ability to bind with the phosphatase MoPpe1, thereby facilitating crosstalk between the TOR pathway and the cell wall integrity pathway [40]. In this study, we found that the TapA and TipA played a crucial role in sclerotia production and cell wall stress in A. flavus (Fig. 5B, S8). The tapA single-copy deletion mutant and the tipA deletion mutant exhibited decreased sensitivity to rapamycin (Fig. 5A). However, the interaction between TapA and TipA is not yet clear in A. flavus. We speculated that the interrelations and functions of TapA and TipA in A. flavus are comparable to their ortholog in other filamentous fungi.

Earlier version: In S. cerevisiae, Tap42 can bind and regulate all PP2A phosphatases, including Pph3, Pph21, Pph22, Sit4 and Ppg1, participating in the TOR signaling pathway in response to rapamycin treatment or nutrient deprivation [35]. Sit4 plays an important role in cell growth and environmental stimuli responses, such as high glucose environment, high temperature and hygromycin [44]. FgSit4 and FgPpg1 are associated with various cellular processes, including mycelial growth, conidiation and DON biosynthesis virulence in F. graminearum [21]. In addition, FgSit4 and FgPpg1 are involved in the regulation of various signaling pathways, such as the TOR signaling pathway, CWI pathway and lipid droplets biogenesis regulation pathway [22]. In this study, we found that the phosphatase SitA and Ppg1 play critical role in vegetative growth, conidiation, pathogenicity, LD biogenesis, sclerotia formation and aflatoxin biosynthesis in A. flavus(Fig. 6A-6C). Additionally, SitA and Ppg1 also affect cell wall integrity stresses response and rapamycin sensitivity. After the deletion of the SitA and Ppg1 genes, the expression level of genes related to cell wall synthesis has decreased significantly, resulting in damage to the cell wall and showing a higher sensitivity to cell wall integrity stress. Western blotting results showed that the phosphorylation level of MAPK kinase Slt2 has increased significantly in the ΔSitA and ΔPpg1 strains(Fig. 8D), indicating that SitA and Ppg1 have a negative regulation on the CWI pathway. We concluded that phosphatases SitA and Ppg1 serve as a mediator for the TOR and CWI pathways in A. flavus. In M. oryzae, MoPpe1 is the homologue of S. cerevisiae serine/threonine protein phosphatase Sit4/Ppe1, which is required for vegetative growth, conidiation and full virulence. MoPpe1 and MoSap1 function as an adaptor complex linking CWI and TOR signaling and the activation of the TOR pathway leads to suppression of CWI signaling[8]. In addition, the phosphatase MoPpe1 interacts with phosphatase-associated protein MoTip41 to mediate crosstalk between TOR and CWI signaling[8]. These results suggest that phosphatase ppe1 interacts with different proteins to coordinate the TOR and CWI signaling pathways in regulating the growth and pathogenicity of rice blast fungus M. oryzae. Taken together, we speculated that phosphatase SitA and Ppg1 may regulate CWI and TOR pathways differently between species to perform their function.

Revised version: In S. cerevisiae, Tap42 can bind and regulate various PP2A phosphatases, including Pph3, Pph21, Pph22, Sit4, and Ppg1. This interaction plays a crucial role in the TOR signaling pathway, particularly in response to rapamycin treatment or nutrient deprivation [48]. The phosphatase Sit4 plays a crucial role in regulating cell growth, morphogenesis, and virulence in C. albicans [62]. FgSit4 and FgPpg1 are associated with various cellular processes, including mycelial growth, conidiation, DON biosynthesis, and virulence in F. graminearum [23]. In addition, FgSit4 and FgPpg1 are involved in the regulation of various signaling pathways, such as the CWI pathway and lipid droplets biogenesis regulation pathway [23]. In this study, we found that the phosphatase SitA and Ppg1 play critical role in vegetative growth, conidiation, sclerotia formation, and aflatoxin biosynthesis in A. flavus (Fig.6A, 6B, 6C). Additionally, the phosphatase SitA and Ppg1 also affect pathogenicity and LD biogenesis (Fig. 7A, 9C, 9D). In addition, TapA/TipA and SitA/Ppg1 exhibits similar functions in sclerotia production and cell wall stress response. Based on these observations, we hypothesized that TapA-phosphatase complexes play a collaborative role in the regulation of the sclerotia formation and cell wall stress response in A. flavus.

Added in the revised version: After the deletion of the sitA and ppg1 genes, the expression level of genes related to cell wall synthesis has decreased significantly, resulting in damage to the cell wall and showing a higher sensitivity to cell wall integrity stress (Fig. 8C). Western blotting results showed that the phosphorylation level of MAPK kinase Slt2 has increased significantly in the ΔsitA and Δppg1 strains under cell wall stress, indicating that SitA and Ppg1 exert a negative regulatory effect on the CWI pathway (Fig. 8D). In M. oryzae, MoPpe1 and MoSap1 serve as an adaptor complex that connects the CWI and TOR signaling pathways. Activation of the TOR pathway results in the inhibition of the CWI pathway [39]. In addition, the phosphatase MoPpe1 interacts with the phosphatase-associated protein MoTip41, facilitating the exchange of information between the TOR and CWI signaling pathways [39]. These results suggest that the phosphatase Ppe1 engages in interactions with various proteins to effectively coordinate the TOR and CWI signaling pathways, thereby regulating the growth and pathogenicity of the rice blast fungus M. oryzae [39]. Taken together, we speculated that the phosphatase SitA and Ppg1 might exhibit differential regulation in the CWI and TOR pathways across species in order to fulfill their respective functions.

Earlier version: In F. graminearum, the phosphatase complex FgNem1/Spo7-FgPah1 cascade plays important roles in fungal development, DON production and LD biogenesis[22]. In this study, our results indicated that Nem1 and Spo7 are important for asexual development and secondary metabolism of A. flavus. We also found that LD biogenesis mediated by the Nem1/Spo7 and SitA/Ppg1 in A. flavus. These observations imply the PP2A phosphatases are core phosphatase in the LD biogenesis signaling network, there should be other kinases involved in regulating the LD biogenesis. It seems that phosphatases may be related to cross-talking among the TOR and LD biogenesis signaling network. The exact molecular mechanism is required to investigate the relationship between phosphatase and other kinase and to explore regulatory mechanism and biological functions of the LDs.

Revised version: In F. graminearum, the phosphatase complex FgNem1/Spo7-FgPah1 cascade plays significant roles in fungal development, DON production, and lipid droplet (LD) biogenesis [24]. In this study, our results indicated that Nem1 and Spo7 are important for vegetative growth and secondary metabolism of A. flavus (Fig. 10A, 10B). We have also discovered that LD biogenesis is facilitated by the Nem1/Spo7 and SitA/Ppg1 pathways (Fig. 9C, 9D, 10E). These observations suggest that the PP2A phosphatases play a central role as core phosphatases in the LD biogenesis signaling network. However, it is likely that there are other kinases involved in the regulation of LD biogenesis. To gain a comprehensive understanding of the precise molecular mechanism, further research should focus on investigating the interplay between phosphatases and other kinases.

Earlier version: In conclusion, we identified the complex regulatory network that elucidates how crosstalk between TOR, HOG and CWI signaling modulates the development, aflatoxin biosynthesis and pathogenicity of A. flavus. By resolving the functions of genes related to the TOR signaling pathway, our findings can lay a solid foundation for an in-depth understanding of the molecular mechanism of A. flavus pathogenicity, which would provide novel insight for controlling AF contamination and reducing A. flavus infection.

Revised version: In conclusion, we identified the complex regulatory network that provides insight into the interplay between the TOR, HOG, and CWI signaling pathways, and how they collectively regulate the development, aflatoxin biosynthesis, and pathogenicity of A. flavus. By elucidating the functions of genes associated with the TOR signaling pathway, our research could establish a solid foundation for comprehending the molecular mechanism underlying the aflatoxin biosynthesis of A. flavus. This, in turn, can provide novel insights into the management of aflatoxin contamination and the mitigation of A. flavus infection.

Materials and methods

Strains and cultural conditions

Earlier version: The A. flavus strains in this study were shown in S1 Table. All strains were cultured on potato dextrose agar (PDA), yeast extract-sucrose agar (YES) and yeast extract-glucose agar (YGT) at 37℃ for growth and conidiation assays respectively, and YES liquid medium and Potato dextrose broth (PDB) were used to detect aflatoxin production at 29℃.

Revised version: The A. flavus strains in this study were shown in Table S3. All strains were cultivated on three different types of agar media: potato dextrose agar (PDA), yeast extract-sucrose agar (YES), and yeast extract-glucose agar (YGT). The growth and conidiation assays were conducted at 37℃. Additionally, aflatoxin production was assessed using YES liquid medium at 29℃.

Earlier version: Constructions of gene deletion and complementation mutants

Revised version: Constructions of gene deletion and complemented mutants

Added in the revised version: The construction of gene deletion and complemented strains was carried out using the SOE-PCR strategy, following previously established protocols [63]. The primers utilized for the amplification of the flanking sequences of each gene can be found in Table S3. The putative gene deletion mutants were validated through PCR assays using the appropriate primers (Fig. S2), and their confirmation was subsequently obtained through sequencing analysis.

Added in the revised version: Detection of tor and tapA genes copy number in strains

Added in the revised version: Absolute quantification PCR has been used as a common method for molecular detection of gene copy numbers. According to previous reports, the sumo gene is a single-copy gene in A. flavus[64]. Following previously established protocols[65], the primers were used to design the amplification of the sumo reference gene and tor and tapA gene. Using sumo gene as reference, the tor and tapA gene copy number was calculated by standard curve.

Removed in the revised version: The gene deletion and complementation strains were constructed using the SOE-PCR strategy according to previously described methods [32]. The primers used to amplify the flanking sequences of each gene were listed in S2 Table. The putative gene deletion mutants were verified by PCR and RT-PCR assays with the relevant primers (Fig S1) and further confirmed by sequencing analysis.

Growth, conidiation, and sclerotia analysis

Earlier version: To assess mycelial growth and conidia formation, 1 μL of 106 spores/mL conidia of all strains were point inoculated onto the center of YES, PDA and YGT agar medium respectively, and then cultured at 37℃ for 5 days in the dark. The colony diameters were measured every day, and the conidia were washed by 0.07% Tween-20 and counted using a hemocytometer and microscope[37]. For sclerotia formation analysis, each strain was inoculated and grown on CM agar medium at 37℃ in the dark for 7 days, and then 75% ethanol was used to wash away mycelia and conidia on the surface of the medium[60].The assays were repeated at least three times.

Revised version: To evaluate mycelial growth and conidia formation, 1 μL of 106 spores/mL conidia of all strains was point inoculated onto the center of PDA medium, and then cultured at 37℃ for 5 days in the dark. The conidia were washed by 0.07% Tween-20 and counted using a hemocytometer and microscope [43]. For sclerotia formation analysis, each strain was inoculated and cultivated on CM agar medium at 37℃ in the dark for 7 days, and then 75% ethanol was used to wash away mycelia and conidia on the surface of the medium [64]. These assays were repeated at least three times.

Extraction and quantitative analysis of aflatoxin

Earlier version: To extract AF and quantify the AF production, 1 μL of 106 spores/mL conidia of all strains were point inoculated onto liquid YES or PDB medium and cultures were incubated at 29℃ in the dark for 6 days. AF were extracted from the media and detected by thin layer chromatography (TLC)[60].In addition, the standard AFB1 was used for visual comparison, and AF was quantitative analyzed by the Image J tool. The assays were repeated at least three times.

Revised version: To extract and quantify aflatoxins (AFs) production, 1 μL of 106 spores/mL conidia of all strains was point inoculated onto liquid YES medium and cultures were then incubated at 29℃ in the dark for 6 days. AFs were extracted from the media and detected using thin layer chromatography (TLC) [65]. In addition, the standard AFB1 was employed for visual comparison, and AFs were quantitatively analyzed using the Image J tool. The assays were repeated at least three times.

Stress assay

Earlier version: To determine the roles of the gene in A. flavus responsing to various environmental stresses, 1 μL of 106 spores/mL conidia of all strains were point inoculated onto these PDA solid plates supplemented with different agents, including 1M NaCl,1M KCl and 1M CaCl2 treatment for osmotic stress, 100µg/mL SDS, 200 µg/mL Calcofluor white (CFW) and 200 µg/mL Congo red (CR) treatment for cell wall stress. After 5 days of incubation at 37℃ in the dark, the relative inhibition rates were calculated[37]. The assays were repeated at least three times.

Revised version: To determine the roles of the gene in A. flavus responsing to various environmental stresses, 1 μL of 106 spores/mL conidia of all strains were point inoculated onto these PDA solid plates supplemented with various agents. These agents included 1M NaCl, 1M KCl, and 1M CaCl2 for inducing osmotic stress, as well as 100 µg/mL SDS, 200 µg/mL Calcofluor white (CFW), and 200 µg/mL Congo red (CR) for inducing cell wall stress. After 5 days of incubation at 37℃ in the dark, the relative rates of inhibition were determined [43]. The assays were repeated at least three times.

Seeds infection assays

Earlier version: To measure the pathogenicity of A. flavus, all peanut and maize seeds were washed in 0.05% sodium hypochlorite, 75% ethanol and sterile water respectively, and then all peanut and maize seeds were soaked in sterile water supplemented with 106 spores/mL of all strains at 29℃ in the dark for 60 min. After the conidia adhesion, the seeds were cultured at 29℃ for 6 days in the dark[60].

Revised version: To measure the pathogenicity of A. flavus, all peanut and maize seeds were washed in 0.05% sodium hypochlorite, 75% ethanol, and sterile water, respectively. Subsequently, all peanut and maize seeds were immersed in sterile water containing 106 spores/mL of all strains at 29℃ in the dark for 60 min. After the conidia adhesion, the seeds were cultured at 29℃ for 6 days in the dark [64].

Quantitative RT-PCR

Earlier version: For real-time quantitative PCR (qPCR) assays, the A. flavus strains were cultured in PDA medium at 29℃ for 3 days in the dark. Total RNA was extracted from the mycelia of all strains, and cDNA was synthesized with First Strand cDNA Synthesis Kit, and then the cDNA was used as the template for qRT-PCR analysis with SYBR Green qPCR mix. The relative transcript level of each gene was calculated by the 2-ΔΔCTmethod[62].

Revised version: For the real-time quantitative PCR (qPCR) assays, the A. flavus strains were cultivated in PDA medium at 29℃ for 3 days in the dark. Total RNA was extracted from the mycelia of all strains, and cDNA was synthesized with the First Strand cDNA Synthesis Kit, and then the cDNA was used as the template for qRT-PCR analysis with SYBR Green qPCR mix. The relative transcript level of each gene was calculated by the 2-ΔΔCTmethod [66].

Western blot analysis

Earlier version: To extract the total protein, the mycelia of each strain were cultured in YES liquid medium for 48 h, then the sample was grinded under liquid nitrogen to a fine powder and transfered to a 1.5 mL microcentrifuge tube added with RIPA lysis buffer and 1 mM PMSF protease inhibitor. Proteins were separated by 12% SDS-PAGE and transferred onto a polyvinylidene fluoride (PVDF) membrane. Phosphorylation levels of Hog1 and Slt2 were detected by Phospho-p38 MAPK antibody (Cell Signaling Technology, Boston, MA, USA) and Phospho-p44/42 MAPK (Erk 1/2) antibody (Cell signaling Technology, MA, USA), respectively[60]. The experiment was conducted three times independently.

Revised version: To extract the total protein, the mycelia of each strain were cultivated in YES liquid medium for 48 h, then the sample was grinded under liquid nitrogen to a fine powder and transferred to a 1.5 mL microcentrifuge tube added with RIPA lysis buffer and 1 mM PMSF protease inhibitor. Proteins were separated by 12% SDS-PAGE and transferred onto a polyvinylidene fluoride (PVDF) membrane. Phosphorylation levels of Hog1 and Slt2 were detected by Phospho-p38 MAPK antibody (Cell Signaling Technology, Boston, MA, USA) and Phospho-p44/42 MAPK (Erk 1/2) antibody (Cell signaling Technology, MA, USA), respectively. The experiment was conducted three times independently.

Statistical analysis

Earlier version: All data were measured with means ± standard deviation (SD) of three biological replicates. Graph Pad Prism 7 software was used for statistical and significance analysis, and p-values < 0.05 was represented statistical significance. Student’s t-test was used to compare the two means for differences, whereas Tukey’s multiple comparisons were used for testing multiple comparisons.

Revised version: All data were measured with means ± standard deviation (SD) of three biological replicates. Graph Pad Prism 7 software was used for conducting statistical and significance analysis, and p-values < 0.05 was represented statistical significance. Student’s t-test was used to compare the two means for differences, whereas Tukey’s multiple comparisons were used for testing multiple comparisons.

Supporting information